```python
# Initialize Otter
import otter
grader = otter.Notebook("lab09.ipynb")
```

# Lab 09: Regression

Welcome to Lab 9.

Today we will get some hands-on practice with linear regression. You can find more information about this topic in the textbook section entitled [The Regression Line](https://www.inferentialthinking.com/chapters/15/2/Regression_Line.html#the-regression-line).

```python
# Run this cell, but please don't change it.

# These lines import the Numpy and Datascience modules.
import numpy as np
from datascience import *

# These lines do some fancy plotting.
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plots
plots.style.use('fivethirtyeight')
```

## 1. How Faithful is Old Faithful? 

Old Faithful is a geyser in Yellowstone National Park that is famous for eruption on a fairly regular schedule. Run the cell below to see Old Faithful in action.

```python
# For the curious: this is how to display a YouTube video in a
# Jupyter notebook.  The argument to YouTubeVideo is the part
# of the URL (called a "query parameter") that identifies the
# video.  For example, the full URL for this video is:
#   https://www.youtube.com/watch?v=wE8NDuzt8eg
from IPython.display import YouTubeVideo
YouTubeVideo("wE8NDuzt8eg")
```

Some of Old Faithful's eruptions last longer than others.  Whenever there is a long eruption, it is usually followed by an even longer wait before the next eruption. If you visit Yellowstone, you might want to predict when the next eruption will happen, so that you can see the rest of the park instead of waiting by the geyser.
 
Today, we will use a dataset on eruption durations and waiting times to see if we can make such predictions accurately with linear regression.

The dataset has one row for each observed eruption.  It includes the following columns:
- `duration`: Eruption duration, in minutes
- `wait`: Time between this eruption and the next, also in minutes

Run the next cell to load the dataset.

```python
faithful = Table.read_table('faithful.csv')
faithful
```

We would like to use linear regression to make predictions, but that won't work well if the data aren't roughly linearly related.  To check that, we should look at the data.

<!-- BEGIN QUESTION -->

### Question 1.1.

Make a scatter plot of the data.  It's conventional to put the column we want to predict on the vertical axis and the other column on the horizontal axis.


```python
...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 1.2.
Are eruption duration and waiting time roughly linearly related based on the scatter plot above? How could you describe the relationship between duration and wait time? Be specific and use appropriate vocabulary in your response.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

We're going to continue with the assumption that they are linearly related, so it's reasonable to use linear regression to analyze this data.

We'd next like to plot the data in standard units. If you don't remember the definition of standard units, reviewing the textbook section entitled [Variability](https://www.inferentialthinking.com/chapters/14/2/Variability.html#standard-units) might help.

### Question 1.3.
Compute the mean and standard deviation of the eruption durations and waiting times and assign them to `duration_mean`, `duration_std`, `wait_mean`, and `wait_std` respectively. Then, use these values to compute the arrays `duration_su` and `wait_su` which contain the eruption durations and waiting times in standard units.

The provided code will create a Table called `faithful_standard` containing the eruption durations and waiting times in standard units based off of the values in `duration_su` and `wait_su`. The columns will be named `duration (standard units)` and `wait (standard units)`.

```python
duration_mean = ...
duration_std = ...
wait_mean = ...
wait_std = ...
duration_su = ...
wait_su = ...

faithful_standard = Table().with_columns(
    "duration (standard units)", duration_su
    "wait (standard units)", wait_su
faithful_standard
```

```python
grader.check("q1_3")
```

<!-- BEGIN QUESTION -->

### Question 1.4.
Plot the data again, but this time in standard units.


```python
...
```

<!-- END QUESTION -->

You'll notice that this plot looks the same as the last one!  However, the data and axes are scaled differently.  It's always important to read the ticks on the axes to ensure you know what units you are working in.

### Question 1.5.

Estimate the value of the correlation coefficient `r` using the scatterplot above. Then, compute the exact value and assign it to the variable `r`.

**Hint:** Use the `faithful_standard` table in your calculation. The section in the textbook entitled [Correlation](https://www.inferentialthinking.com/chapters/15/1/Correlation.html#calculating-r) could be helpful if you need some guidance.

```python
r = ...
r
```

```python
grader.check("q1_5")
```

## 2. The Regression Line
Recall that the value of the correlation coefficient is equal to the slope of the regression line when the data are put in standard units.

The next cell uses a custom function to plot the regression line:

$$\text{waiting time in standard units} = r \times \text{eruption duration in standard units}.$$

in standard units using your computed value for `r` as the slope. Then, it plots the data in standard units again, for comparison.

```python
def plot_data_and_line(dataset, x, y, point_0, point_1):
    """Makes a scatter plot of the dataset, along with a line passing through two points."""
    dataset.scatter(x, y, label="data")
    xs, ys = zip(point_0, point_1)
    plots.plot(xs, ys, label="regression line")
    plots.legend(bbox_to_anchor=(1.5,.8))

plot_data_and_line(faithful_standard, 
                   "duration (standard units)", 
                   "wait (standard units)", 
                   [-2, -2*r], 
                   [2, 2*r])
```

### Converting standard units back to original units

Recall that in order to convert standard units back to the original units, you'd have to "stretch" (or "compress") the horizontal values by `duration_std` and its vertical values by `wait_std`.

### Question 2.1.
Calculate the slope of the regression line in original units, and assign it to `slope`.

If the "stretching" explanation above is not intuitive, consult the section in the textbook entitled [The Regression Line](https://www.inferentialthinking.com/chapters/15/2/Regression_Line.html#the-equation-of-the-regression-line).


```python
slope = ...
slope
```

```python
grader.check("q2_1")
```

Since we know that the regression line passes through (0, 0) in standard units, the regression line also passes through the point `(duration_mean, wait_mean)` in the original units. The equation for the regression line in point-slope form would be:

$$\text{waiting time} - \verb|wait_mean| = \texttt{slope} \times (\text{eruption duration} - \verb|duration_mean|)$$

### Question 2.2.
Calculate the intercept in original units and assign it to `intercept`. It might help to rearrange the equation for the regression line into slope-intercept form to determine the expression for computing the intercept.

```python
intercept = ...
intercept
```

```python
grader.check("q2_2")
```

## 3. Investigating the regression line
The slope and intercept tell you exactly what the regression line looks like.  To predict the waiting time for an eruption, multiply the eruption's duration by `slope` and then add `intercept`.

### Question 3.1.
Use the slope and intercept of the regression line to compute the predicted waiting time for:
* an eruption that lasts 2 minutes, and
* an eruption that lasts 5 minutes

Assign the predicted values to `two_minute_predicted_waiting_time` and `five_minute_predicted_waiting_time` respectively.

```python
two_minute_predicted_waiting_time = ...
five_minute_predicted_waiting_time = ...

# Here is a helper function to print out your predictions.
# Don't modify the code below.

def print_prediction(duration, predicted_waiting_time):
    print("After an eruption lasting", duration, "minutes, we predict you'll wait approximately", round(predicted_waiting_time), "minutes until the next eruption.")
    
print_prediction(2, two_minute_predicted_waiting_time)
print_prediction(5, five_minute_predicted_waiting_time)
```

```python
grader.check("q3_1")
```

The next cell plots the line that goes between the two predicted points, which is (a segment of) the regression line.

```python
plot_data_and_line(faithful, "duration", "wait", 
                   [2, two_minute_predicted_waiting_time], 
                   [5, five_minute_predicted_waiting_time])
```

### Question 3.2.
Now, make predictions for the waiting time after *each* eruption in the `faithful` table. Of course, you know exactly what the actual waiting times were. We are doing this so we can see how accurate the regression line predictions are. 

Assign the predicted values in a column named `"predicted wait"` in a new table called `faithful_predictions`.  

Its first row should look like:

|duration|wait|predicted wait|
|-|-|-|
|3.6|79|72.1011|

**Hint:** Your answer can be just one line.  There is no need for a `for` loop; use array arithmetic instead.

```python
faithful_predictions = ...
faithful_predictions
```

```python
grader.check("q3_2")
```

### Question 3.3.
To investigate how well your predictions compare to what actually happened, compute the **residual** (or error) for each eruption in the dataset.  The **residual** is computed by subtracting the predicted waiting time from the actual waiting time. Add the residual value for each eruption to the `faithful_predictions` Table as a new column called `"residual"` and name the resulting new Table `faithful_residuals`.

**Hint:** Your code will be much simpler (and faster!) if you don't use a `for` loop but instead use numpy/array operations.

```python
faithful_residuals = ...
faithful_residuals
```

```python
grader.check("q3_3")
```

Here is a plot of the residuals you computed.  Each point corresponds to the calculated error of one eruption.  The residual plot shows how the predictions over- or under-estimated the actual waiting time.

```python
faithful_residuals.scatter('duration', 'residual', color = 'r')
```

The plot indicates that there isn't really a pattern in the residuals; the predictions appear to be randomly larger or smaller than the actual value. The result confirms that it was reasonable to use linear regression to create a model to make the predictions.

## 4. How accurate are different predictions?
Earlier, you should have found that the correlation is fairly close to 1, so the line fits fairly well on the training data.  That means the residuals are overall small (close to 0) in comparison to the waiting times.

We can see that visually by plotting the waiting times and residuals together:

```python
faithful_residuals.scatter('duration', 'wait', label = 'actual waiting time', color = 'blue')
plots.scatter(faithful_residuals.column('duration'), faithful_residuals.column('residual'), label = 'residual', color = 'r')
plots.plot([2, 5], [two_minute_predicted_waiting_time, five_minute_predicted_waiting_time], label = 'regression line')
plots.legend(bbox_to_anchor = (1.7, .8));
```

However, unless you have a strong reason to believe that the linear regression model is a true model of the natural phenomenon, which very rarely occurs, you should be wary of applying your prediction model to data that are very different from the training data.

### Question 4.1.
In `faithful`, no eruption lasted exactly 0 minutes, 2.5 minutes, or 1 hour.  Using this line, what is the predicted waiting time for an eruption that lasts 0 minutes?  2.5 minutes?  1 hour? Assign the values for their predictions to `zero_minute_predicted_waiting_time`, `two_point_five_minute_predicted_waiting_time`, and `hour_predicted_waiting_time` respectively.

```python
zero_minute_predicted_waiting_time = ...
two_point_five_minute_predicted_waiting_time = ...
hour_predicted_waiting_time = ...
print_prediction(0, zero_minute_predicted_waiting_time)
print_prediction(2.5, two_point_five_minute_predicted_waiting_time)
print_prediction(60, hour_predicted_waiting_time)
```

```python
grader.check("q4_1")
```

<!-- BEGIN QUESTION -->

### Question 4.2.
For each prediction, state whether you think it's reliable and explain your reasoning as to why or why not. Think about what the input values of 0, 2.5 and 60 would represent in real life.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

# Submitting your work
You're done with this assignment! Assignments should be turned in using the following best practices:
1. Save your notebook.
2. Restart the kernel and run all cells up to this one.
3. Run the cell below with the code `grader.export(...)`. This will re-run all the tests. Make sure they are passing as you expect them to.
4. Download the file named `lab09_<date-time-stamp>.zip`, found in the explorer pane on the left side of the screen. **Note**: Clicking on the link in this notebook may result in an error, it's best to download from the file explorer panel.
5. Upload `lab09_<date-time-stamp>.zip` to the corresponding assignment on Canvas.

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit.

```python
grader.export(pdf=False, force_save=True)
```

Tests for q1_3 (points: None)
```
>>> int(round(wait_std)) == 14 and int(round(duration_std)) == 1
True
```
failure_message: ❌ `wait_std` and `duration_std` are not reasonable values for this data
success_message: ✅ `wait_std` and `duration_std` are reasonable values
```
>>> np.isclose(abs(sum(faithful_standard.column(0))), 0) and np.isclose(abs(sum(faithful_standard.column(1))), 0)
True
```
failure_message: ❌ The `faithful_standard` Table has at least one column that does not sum to 0
success_message: ✅ The `faithful_standard` Table has columns that sum to 0


Tests for q1_5 (points: None)
```
>>> np.isclose(r, 0.9008111683218132)
True
```
failure_message: ❌ `r` is not the correct value
success_message: ✅ `r` is the correct value


Tests for q2_1 (points: None)
```
>>> 9 < slope < 11
True
```
failure_message: ❌ The value for `slope` does not seem reasonable
success_message: ✅ The value for `slope` seems reasonable


Tests for q2_2 (points: None)
```
>>> -31 < intercept < 35
True
```
failure_message: ❌ The value for `slope` does not seem reasonable
success_message: ✅ The value for `slope` seems reasonable


Tests for q3_1 (points: None)
```
>>> 50 < two_minute_predicted_waiting_time < 60
True
```
failure_message: ❌ `two_minute_predicted_waiting_time` is too far off to be correct
success_message: ✅ `two_minute_predicted_waiting_time` is close to, and possibly is, correct
```
>>> 80 < five_minute_predicted_waiting_time < 90
True
```
failure_message: ❌ `five_minute_predicted_waiting_time` is too far off to be correct
success_message: ✅ `five_minute_predicted_waiting_time` is close to, and possibly is, correct


Tests for q3_2 (points: None)
```
>>> set(faithful_predictions.labels) == set(['duration', 'wait', 'predicted wait'])
True
```
failure_message: ❌ The column labels in `faithful_predictions` are not correct
success_message: ✅ The column labels in `faithful_predictions` are correct
```
>>> faithful_predictions.num_rows == 272
True
```
failure_message: ❌ `faithful_predictions` does not have the correct number of rows
success_message: ✅ `faithful_predictions` has the correct number of rows
```
>>> np.isclose(np.mean(faithful_predictions.column('predicted wait')), 70.8970588235294)
True
```
failure_message: ❌ The mean of the 'predicted wait' column is incorrect, indicating that some values are incorrect or imprecise
success_message: ✅ The mean of the 'predicted wait' column is correct, indicating the values are correct
```
>>> np.isclose(np.std(faithful_predictions.column('predicted wait')), 12.223971537522273)
True
```
failure_message: ❌ The standard deviation of the 'predicted wait' column is incorrect, indicating that some values are incorrect or imprecise
success_message: ✅ The standard deviation of the 'predicted wait' column is correct, indicating the values are correct


Tests for q3_3 (points: None)
```
>>> set(faithful_residuals.labels) == set(['duration', 'wait', 'predicted wait', 'residual'])
True
```
failure_message: ❌ The column labels in `faithful_residuals` are not correct
success_message: ✅ The column labels in `faithful_residuals` are correct
```
>>> faithful_residuals.num_rows == 272
True
```
failure_message: ❌ `faithful_residuals` does not have the correct number of rows
success_message: ✅ `faithful_residuals` has the correct number of rows
```
>>> np.isclose(np.mean(faithful_residuals.column('residual')), 0)
True
```
failure_message: ❌ The mean of the 'residual' column is incorrect, indicating that some values are incorrect or imprecise
success_message: ✅ The mean of the 'residual' column is correct, indicating the values are correct
```
>>> np.isclose(np.std(faithful_residuals.column('residual')), 5.892226635893954)
True
```
failure_message: ❌ The standard deviation of the 'residual' column is incorrect, indicating that some values are incorrect or imprecise
success_message: ✅ The standard deviation of the 'residual' column is correct, indicating the values are correct


Tests for q4_1 (points: None)
```
>>> 12 - zero_minute_predicted_waiting_time * 1.4 / 4 <= 0.35
True
```
failure_message: ❌ `zero_minute_predicted_waiting_time` is too far off from the correct answer
success_message: ✅ `zero_minute_predicted_waiting_time` is reasonably close the correct answer
```
>>> 2 - two_point_five_minute_predicted_waiting_time / 35 <= 0.4
True
```
failure_message: ❌ `two_point_five_minute_predicted_waiting_time` is too far off from the correct answer
success_message: ✅ `two_point_five_minute_predicted_waiting_time` is reasonably close the correct answer
```
>>> (26 - hour_predicted_waiting_time / 30) / 10 <= 0.43
True
```
failure_message: ❌ `hour_predicted_waiting_time` is too far off from the correct answer
success_message: ✅ `hour_predicted_waiting_time` is reasonably close the correct answer
