```python
# Initialize Otter
import otter
grader = otter.Notebook("hw02.ipynb")
```

# Homework 02: Arrays and Tables

Please complete this notebook by filling in the cells provided. Before you begin, execute the following cell to load the provided tests.

```python
# Don't change this cell; just run it. 
# When you log-in please hit return (not shift + return) after typing in your email
import numpy as np
from datascience import *
```

**Helpful Resource:**
- [Python Reference](http://data8.org/sp22/python-reference.html): Cheat sheet of helpful array & table methods used in this course.

**Recommended Reading:**
- [What is Data Science](http://www.inferentialthinking.com/chapters/01/what-is-data-science.html)
- [Causality and Experiments](http://www.inferentialthinking.com/chapters/02/causality-and-experiments.html) 
- [Programming in Python](http://www.inferentialthinking.com/chapters/03/programming-in-python.html)

For all problems that you must write explanations and sentences for, you **must** provide your answer in the designated space. Moreover, throughout this homework and all future ones, please be sure to not re-assign variables throughout the notebook! For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Otherwise, you will fail tests that you thought you were passing previously!

**Note: This homework has hidden tests on it. Additional tests will be run once your homework is submitted for grading. While you may pass all the tests you have access to before submission, you may not earn full credit if you do not pass the hidden tests as well.**. 

Many of the tests you have access to before submitting only test to ensure you have given an answer that is formatted correctly and/or you have given an answer that *could* make sense in context. For example, a test you have access to while completing the assignment may check that you selected a valid choice for a multiple choice problem (1, 2, or 3) or that your answer is an integer between 0 and 50 if asked to count a subset of states in the United States. The tests that are run after submission will evaluate your work for accuracy. **Do not assume that just because all your tests pass before submission means that your answers are correct!**

Consult with your teacher and course syllabus for information and policies regarding appropriate collaboration with other students, appropriate use of AI tools, and submission of late work.

## 1. Creating Arrays

#### Question 1.1.
Make an **array** called `weird_numbers` containing the following numbers (in the given order):

1. $-2$
2. $\sin(12)$
3. $3$
4. $5^{\cos(12)}$

**Hint:** `sin` and `cos` are functions in the `math` module. Importing modules is covered in a previous lab! You should import the module in your solution cell below.

*Note:* You may be familiar with Python lists if you've worked with Python before. If you have, please note that Python lists are different and behave differently than `numpy` arrays. In this course, we use `numpy` arrays, so please make an **array**, not a Python list. If you've never heard of Python lists before, feel free to ignore this message!

```python
# Our solution involved one extra line of code before creating
# weird_numbers.
...
weird_numbers = ...
weird_numbers
```

```python
grader.check("q1_1")
```

#### Question 1.2.
Make an array called `book_title_words` containing the following three strings: "Eats", "Shoots", and "and Leaves".

```python
book_title_words = ...
book_title_words
```

```python
grader.check("q1_2")
```

### `join`
All strings in Python have a built-in method called `join`.  The `join` method takes one argument as an input, an array of strings. The join method returns a single string as an output. The output string of `a_string.join(an_array)` is a single string that's the [concatenation](https://en.wikipedia.org/wiki/Concatenation) ("putting together") of all the strings in `an_array`, with `a_string`  inserted in between each string from the array. Run the cell below to see it in action:

```python
area_code = '828'
prefix = '347'
line_number = '9100'

phone_number = '-'.join(make_array(area_code, prefix, line_number))
phone_number
```

Notice how the single dash, `-`, was inserted between each of the three strings contained in the array to generate the phone number for the NCSSM-Morganton front desk.

#### Question 1.3. 
Use the array `book_title_words` from the previous question and the method `join` to make two different strings:

1. "Eats, Shoots, and Leaves" (call this one `with_commas`)
2. "Eats Shoots and Leaves" (call this one `without_commas`)

```python
with_commas = ...
without_commas = ...

print('with commas: ', with_commas)
print('without commas: ', without_commas)
```

```python
grader.check("q1_3")
```

## 2. Indexing Arrays

These exercises give you practice accessing individual elements of arrays.  In Python (and in many programming languages), elements are accessed by *index*, and  the first element is at index `0`.

***Note:* If you have previous coding experience, you may be familiar with bracket notation. DO NOT use bracket notation when indexing (i.e. `arr[0]`), as this can yield different data type outputs than what we will be expecting. This can cause you to fail an autograder test.**

Be sure to refer to the [Python Reference](http://data8.org/sp22/python-reference.html) on the website if you feel stuck!

#### Question 2.1.
The cell below creates an array of some numbers.  Set `third_element` to the third element of `some_numbers` using the `item` method.

```python
some_numbers = make_array(-1, -3, -6, -10, -15)

third_element = ...
third_element
```

```python
grader.check("q2_1")
```

#### Question 2.2.
The next cell creates a table that displays some information about the elements of `some_numbers` and their order.  Run the cell to see the partially-completed table, then fill in the missing information (the cells that say "Ellipsis") by assigning `blank_a`, `blank_b`, `blank_c`, and `blank_d` to the correct elements in the table.

**Hint:** Replace the ... with strings or numbers to make the Table consistent with the array from the previous problem.

```python
blank_a = ...
blank_b = ...
blank_c = ...
blank_d = ...

# Don't change anything below this line!
elements_of_some_numbers = Table().with_columns(
    "English name for position", make_array("first", "second", blank_a, blank_b, "fifth"),
    "Index",                     make_array(blank_c, 1, 2, blank_d, 4),
    "Element",                   some_numbers)
elements_of_some_numbers
```

```python
grader.check("q2_2")
```

#### Question 2.3.
You'll sometimes want to find the *last* element of an array.  Suppose an array has $142$ elements.  What is the *index* of its last element?

```python
index_of_last_element = ...
```

```python
grader.check("q2_3")
```

More often than not, you don't know the number of elements in an array, also known as the array's *length*. For example, it might be a large data set you found on the Internet and aren't too familiar with. The function `len` takes a single argument, an array (or any collection of objects in Python), and returns the `len`gth of that array. The length will always be returned as an integer.

#### Question 2.4.
The cell below loads a data set that contains information about deceased US presidents, starting with George Washington, and assigns it to the Table `presidents_births`. The code then selects a single column from the Table that contains information about the year of birth of each of the presidents and assigns the column data to an array called `president_birth_years`.  Recall that calling `.column(...)` on a table returns an array containing the values of the column specified, in this case the `Birth Year` column of the `president_births` table. 

The last element in that array is the most recent birth year of any deceased US president. Assign that year to `most_recent_birth_year`. Your solution should used both the `len(...)` and `.item(...)` functions to determine the answer, and you should not have to know the number of rows in `president_birth_years` to answer this question. Your code should be able to work no matter how big or small `president_birth_years` might be.

```python
presidents_births = Table.read_table("president_births.csv")
president_birth_years = Table.read_table("president_births.csv").column('Birth Year')

most_recent_birth_year = ...
most_recent_birth_year
```

```python
grader.check("q2_4")
```

#### Question 2.5.
Finally, assign `sum_of_birth_years` to the sum of the first, sixteenth, and last birth year in `president_birth_years`. Your solution should not include the actual years in the calculation, but rather references to the array where those years are stored.

```python
sum_of_birth_years = ...
sum_of_birth_years
```

```python
grader.check("q2_5")
```

## 3. Basic Array Arithmetic

#### Question 3.1.
Multiply the numbers `10`, `-20`, `300,000,000`, and `250` by `157`. Assign each variable below such that `first_product` is assigned to the result of $10 \cdot 157$, `second_product` is assigned to the result of $-20 \cdot 157$, and so on. 

For this question, **don't** use arrays.

```
BEGIN QUESTION
name: q3_1
```

```python
first_product = ...
second_product = ...
third_product = ...
fourth_product = ...
print(first_product, second_product, third_product, fourth_product)
```

```python
grader.check("q3_1")
```

#### Question 3.2.
Now, do the same calculation, but using an array called `numbers` that contains the 4 original numbers (`10`, `-20`, `300,000,000`, and `250`) and only a single multiplication (`*`) operator.  Store the 4 results in an array named `products`.

```python
numbers = ...
products = ...
products
```

```python
grader.check("q3_2")
```

#### Question 3.3.
Oops, we made a typo!  Instead of `157`, we wanted to multiply each number by `1577`. Compute the correct products in the cell below using array arithmetic. The purpose of this question is to help you realize that it's much easier to perform (and in this case, change) arithmetic on an array all at once, rather than having to do so once for each element in the array. Notice that your job is going to be really easy if you previously defined an array containing the 4 numbers.

```python
correct_products = ...
correct_products
```

```python
grader.check("q3_3")
```

#### Question 3.4.
We've loaded an array of temperatures in the next cell.  Each number is the highest temperature observed on a day at a climate observation station, mostly from the US.  Since they're from the US government agency [NOAA](noaa.gov), all the temperatures are in Fahrenheit.  Convert them all to Celsius by first subtracting 32 from them, then multiplying the results by $\frac{5}{9}$. Then, after converting to Celsius **round** the each of the final calculations to the nearest integer using the `np.round` function. `np.round` works a lot like the built-in `round` function in Python, but it can operate on entire arrays.

```python
max_temperatures = Table.read_table("temperatures.csv").column("Daily Max Temperature")

celsius_max_temperatures = ...
celsius_max_temperatures
```

```python
grader.check("q3_4")
```

#### Question 3.5.
The cell below loads all the *lowest* temperatures from each day in Fahrenheit. Each temperature corresponds to the same date and place the max temperatures used in the previous were recorded.

Compute the size of the daily temperature range for each day. That is, compute the difference between each daily maximum temperature and the corresponding daily minimum temperature.  **Pay attention to the units, and give your answer in Celsius!** 

Make sure **NOT** to round your answer for this question! Note: Remember that in the previous question, `celsius_max_temperatures` was rounded, so you probably don't want to use that in this question.

```python
min_temperatures = Table.read_table("temperatures.csv").column("Daily Min Temperature")

celsius_temperature_ranges = ...
celsius_temperature_ranges
```

```python
grader.check("q3_5")
```

## 4. Old Faithful

Old Faithful is a geyser in Yellowstone that erupts every 44 to 125 minutes (according to [Wikipedia](https://en.wikipedia.org/wiki/Old_Faithful)). People are [often told that the geyser erupts every hour](http://yellowstone.net/geysers/old-faithful/), but in fact the waiting time between eruptions is more variable. Let's take a look.

#### Question 4.1.
The first line below assigns `waiting_times` to an array of 272 consecutive waiting times between eruptions, taken from a classic 1938 dataset. Assign the names `shortest`, `longest`, and `average` so that the `print` statement is correct.

```python
waiting_times = Table.read_table('old_faithful.csv').column('waiting')

shortest = ...
longest = ...
average = ...

print("Old Faithful erupts every", shortest, "to", longest, "minutes and every", average, "minutes on average.")
```

```python
grader.check("q4_1")
```

#### Question 4.2
Suppose you're curious about the changes between consecutive wait times for an eruption, specifically when those changes result in a shorter wait time. For example, the third eruption occurred after 74 minutes and the fourth after 62 minutes, so the decrease in waiting time was $62 - 74 = -12$ minutes, for a decrease of 12 minutes. Notice, we report the decrease as a positive number, even though the calculation resulted in a negative value.

First, compute the change in all consecutive wait times, and then assign the biggest decrease in waiting time between two consecutive eruptions to `biggest_decrease`.  You'll want to use an array arithmetic function [mentioned in the textbook](https://inferentialthinking.com/chapters/05/1/Arrays.html#functions-on-arrays) called `np.diff`. The illustration below gives you a quick idea on how it works:

![](array_diff.png)

```python
differences = ...
biggest_decrease = ...
biggest_decrease
```

```python
grader.check("q4_2")
```

#### Question 4.3.
The `faithful_with_eruption_nums` table contains two columns: `eruption_number`, which represents the number of that eruption, and `waiting`, which represents the time spent waiting after that eruption. 

For example, take the first two rows of the table:

```python
faithful = Table.read_table('old_faithful.csv').drop("eruptions")
faithful_with_eruption_nums = faithful.with_column("eruption number", np.arange(faithful.num_rows) + 1).select(1, 0)
faithful_with_eruption_nums.show(2)
```

We can read this as follows: after the first eruption, we waited 79 minutes for the second eruption. Then, after the second eruption, we waited 54 minutes for the third eruption.

Suppose Oscar and Wendy started watching Old Faithful at the start of the first eruption and they watch until the end of the tenth eruption. How many minutes will they spend waiting for eruptions?

Use Table operations (like `take` or `where`) to create a subset of the Table provided and assign the table to `first_nine_waiting_times`. Then, use array operations on a column of the Table to compute and assign your answer to `total_waiting_time_until_tenth`.

```python
first_nine_waiting_times = ...
total_waiting_time_until_tenth = ...
total_waiting_time_until_tenth
```

```python
grader.check("q4_3")
```

#### Question 4.4.
Let’s imagine your guess for the next wait time was always just the length of the previous waiting time. If you always guessed the previous waiting time, how big would your error in guessing the waiting times be, on average? Assign your answer to `average_error`.

For example, the first four waiting times are 79, 54, 74, and 62. After the first eruption you guess the next eruption will happen after 79 minutes, but it occurs after 54 resulting in an error of $\lvert 79 - 54 \rvert$. For the next eruption you now guess it will take 54 minutes, but it takes 74, so the error is $\lvert 54 - 74 \rvert$., and so on. Therefore, the average difference between your guess and the actual time for just the second, third, and fourth eruptions would be $\frac{\lvert 79-54 \rvert+ \lvert54-74\rvert + \lvert 74-62 \rvert}{3} = 19$. 

**Hint:** Use operations to break the process into a series of calculations performed on the entire array.

```python
average_error = ...
average_error
```

```python
grader.check("q4_4")
```

## 5. Tables


#### Question 5.1.
Suppose you have 4 apples, 3 oranges, and 3 pineapples.  (Perhaps you're using Python to solve a high school Algebra problem.)  Create a table that contains this information.  It should have two columns: `fruit name` and `count`.  Assign the new table to the variable `fruits`.

**Note:** Use lower-case and singular words for the name of each fruit, like `"apple"`.

Your table should look like:

| fruit name | count |
|------------|-------|
| apple      | 4     |
| orange     | 3     |
| pineapple  | 3     |

```python
# Our solution uses 1 statement split over 3 lines. You can write yours on 1 line if you wish.
fruits = ...
    ...
    ...
fruits
```

```python
grader.check("q5_1")
```

#### Question 5.2.
The file `inventory.csv` can be found in the same folder as this notebook and contains information about the inventory at a fruit stand.  Each row represents the contents of one box of fruit. Load it as a Table named `inventory` using the `Table.read_table()` function. `Table.read_table(...)` takes one argument (data file name in string format) and returns a table.

```python
inventory = ...
inventory
```

```python
grader.check("q5_2")
```

#### Question 5.3.
The file `sales.csv` contains the number of fruit sold from each box last Saturday.  It has an extra column called "price per fruit (\$)" that's the price *per item of fruit* for fruit in that box.  The rows are in the same order as the `inventory` table.  

Load these data into a table called `sales`.

```python
sales = ...
sales
```

```python
grader.check("q5_3")
```

#### Question 5.4.
How many fruits did the store sell in total on that day?

```python
total_fruits_sold = ...
total_fruits_sold
```

```python
grader.check("q5_4")
```

#### Question 5.5.
What was the store's total revenue (the total price of all fruits sold) on that day?

**Hint:** If you're stuck, think first about how you would compute the total revenue from just the grape sales.

```python
total_revenue = ...
total_revenue
```

```python
grader.check("q5_5")
```

#### Question 5.6.
Make a new table called `remaining_inventory`.  It should have the same rows and columns as `inventory`, except that the amount of fruit sold from each box should be subtracted from that box's **original** count, so that the `count` column is **updated to be** the amount of fruit remaining after Saturday. 

The provided solution spans multiple lines, but you can write yours on a single line if you prefer.

```python
remaining_inventory = ...
    ...
    ...
    ...
remaining_inventory
```

```python
grader.check("q5_6")
```

# Submitting your work
You're done with this assignment! Assignments should be turned in using the following best practices:
1. Save your notebook.
2. Restart the kernel and run all cells up to this one.
3. Run the cell below with the code `grader.export(...)`. This will re-run all the tests. Make sure they are passing as you expect them to.
4. Download the file named `hw02_<date-time-stamp>.zip`, found in the explorer pane on the left side of the screen. **Note**: Clicking on the link in this notebook may result in an error, it's best to download from the file explorer panel.
5. Upload `hw02_<date-time-stamp>.zip` to the corresponding assignment on Canvas.

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit.

```python
grader.export(pdf=False, force_save=True)
```

Tests for q1_1 (points: None)
```
>>> isinstance(weird_numbers, np.ndarray)
True
```
failure_message: ❌ You did not create an array. Perhaps you created a list or another type of collection?
success_message: ✅ You successfully created an array.
```
>>> len(weird_numbers)
4
```
failure_message: ❌ Your array does not have the correct number of elements in it.
success_message: ✅ Your array has the correct number of elements in it.


Tests for q1_2 (points: None)
```
>>> type(book_title_words) == np.ndarray
True
```
failure_message: ❌ You did not create an array. Perhaps you created a list or another type of collection?
success_message: ✅ You successfully created an array.
```
>>> len(book_title_words)
3
```
failure_message: ❌ Your array does not have the correct number of elements in it
success_message: ✅ Your array has the correct number of elements in it
```
>>> not any([',' in text for text in book_title_words])
True
```
failure_message: ❌ You included a comma in one of your string. You should not include the comma.
```
>>> 'and ' in book_title_words.item(2)
True
```
failure_message: ❌ You did not include the word and in any of the strings. Make sure the last string is "and Leaves", not just "Leaves"


Tests for q1_3 (points: None)
```
>>> ',' in with_commas
True
```
failure_message: ❌ The string `with_commas` does not have any commas in it
success_message: ✅ The string `with_commas` has at least one comma in it
```
>>> ',' not in without_commas
True
```
failure_message: ❌ The string `without_commas` has at least one comma in it
success_message: ✅ The string `without_commas` does not have any commas in it
```
>>> len(with_commas)
24
```
failure_message: ❌ The string `with_commas` is not the correct length
success_message: ✅ The string `with_commas` is the correct length
```
>>> len(without_commas)
22
```
failure_message: ❌ The string `without_commas` is not the correct length
success_message: ✅ The string `without_commas` is the correct length


Tests for q2_1 (points: None)
```
>>> third_element in some_numbers
True
```
failure_message: ❌ The number you assigned to `third_element` is not from the array `some_numbers`
success_message: ✅ The number you assigned to `third_element` is from the array `some_numbers`
```
>>> third_element != -10
True
```
failure_message: ❌ You incorrectly selected the element at index 3. This is the fourth element of the array.
success_message: ✅ You successfully did not select the element at *index* 3 (this is a common mistake)


Tests for q2_2 (points: None)
```
>>> elements_of_some_numbers.column(0).item(2).lower()
'third'
```
failure_message: ❌ `blank_a` is incorrect. Check your answer and the spelling. There should be no Punctuation or spaces.
success_message: ✅ `blank_a` is correct
```
>>> elements_of_some_numbers.column(0).item(3).lower()
'fourth'
```
failure_message: ❌ `blank_b` is incorrect. Check your answer and the spelling. There should be no Punctuation or spaces.
success_message: ✅ `blank_b` is correct
```
>>> elements_of_some_numbers.column(1).item(0)
0
```
failure_message: ❌ `blank_c` is incorrect. Check your answer and the spelling. There should be no Punctuation or spaces.
success_message: ✅ `blank_c` is correct
```
>>> elements_of_some_numbers.column(1).item(3)
3
```
failure_message: ❌ `blank_d` is incorrect. Check your answer and the spelling. There should be no Punctuation or spaces.
success_message: ✅ `blank_d` is correct


Tests for q2_3 (points: None)
```
>>> index_of_last_element
141
```
failure_message: ❌ You incorrectly determined the index. Don't forget, indexing starts at 0.
success_message: ✅ You correctly determined the index


Tests for q2_4 (points: None)
```
>>> type(most_recent_birth_year) == int and most_recent_birth_year > 1900
True
```
failure_message: ❌ Your solution should result in an integer greater than 1900 since it represents a year of birth for the most recently deceased president
success_message: ✅ Your solution is an integer greater than 1900, so it could represent the correct answer


Tests for q2_5 (points: None)
```
>>> sum_of_birth_years > 5100 and sum_of_birth_years < 6000
True
```
failure_message: ❌ The sum of the three years is not greater than 5100 and less than 6000.
success_message: ✅ The sum of the three years is greater than 5100 and less than 6000.


Tests for q3_1 (points: None)
```
>>> first_product
1570
```
failure_message: ❌ You incorrectly computed `first_product`. Double check you've typed the numbers correctly!
success_message: ✅ You correctly computed `first_product`.
```
>>> second_product
-3140
```
failure_message: ❌ You incorrectly computed `second_product`. Double check you've typed the numbers correctly!
success_message: ✅ You correctly computed `second_product`.
```
>>> third_product
47100000000
```
failure_message: ❌ You incorrectly computed `third_product`. Double check you've typed the numbers correctly!
success_message: ✅ You correctly computed `third_product`.
```
>>> fourth_product
39250
```
failure_message: ❌ You incorrectly computed `fourth_product`. Double check you've typed the numbers correctly!
success_message: ✅ You correctly computed `fourth_product`.


Tests for q3_2 (points: None)
```
>>> type(products) == np.ndarray
True
```
failure_message: ❌ You did not create an array. Perhaps you created a list or another type of collection?
success_message: ✅ You successfully created an array.
```
>>> len(products)
4
```
failure_message: ❌ Your array does not have the correct number of elements in it.
success_message: ✅ Your array has the correct number of elements in it.


Tests for q3_3 (points: None)
```
>>> np.allclose(correct_products, np.array([15770, -31540, 473100000000, 394250]))
True
```


Tests for q3_4 (points: None)
```
>>> type(celsius_max_temperatures) == np.ndarray
True
```
failure_message: ❌ You did not create an array. Perhaps you created a list or another type of collection?
success_message: ✅ You successfully created an array.
```
>>> len(celsius_max_temperatures)
65000
```
failure_message: ❌ Your array does not have the correct number of elements in it.
success_message: ✅ Your array has the correct number of elements in it.
```
>>> np.sum(celsius_max_temperatures) != 356705.0
True
```
failure_message: ❌ Check your order of operations. It looks like you multiplied by 5/9, then subtracted 32


Tests for q3_5 (points: None)
```
>>> type(celsius_temperature_ranges) == np.ndarray
True
```
failure_message: ❌ You did not create an array. Perhaps you created a list or another type of collection?
success_message: ✅ You successfully created an array.
```
>>> len(celsius_temperature_ranges)
65000
```
failure_message: ❌ Your array does not have the correct number of elements in it.
success_message: ✅ Your array has the correct number of elements in it.


Tests for q4_1 (points: None)
```
>>> shortest <= average <= longest
True
```
failure_message: ❌ Your `shortest` is not the smallest value and `longest` is not the largest value.
success_message: ✅ Shortest <= Average <= Longest


Tests for q4_2 (points: None)
```
>>> biggest_decrease < 0
False
```
failure_message: ❌ You incorrectly provided an answer that was negative. Remember, we report the decrease as a positive number.
success_message: ✅ You correctly provided an answer that was positive
```
>>> biggest_decrease != 47
True
```
failure_message: ❌ You computed the biggest increase in wait times instead of the biggest decrease
success_message: ✅ You (correctly) did not compute the biggest increase in wait time
```
>>> 30 <= biggest_decrease < 47
True
```
failure_message: ❌ The biggest decrease is between 30 and 47, your calculation is not.
success_message: ✅ The biggest decrease is between 30 and 47
```
>>> biggest_decrease
45
```


Tests for q4_3 (points: None)
```
>>> type(first_nine_waiting_times) == Table
True
```
failure_message: ❌ `first_nine_waiting_times` should be a Table
success_message: ✅ `first_nine_waiting_times` is a Table
```
>>> first_nine_waiting_times.num_rows
9
```
failure_message: ❌ `first_nine_waiting_times` does not have the correct number of rows
success_message: ✅ `first_nine_waiting_times` has the correct number of rows


Tests for q4_4 (points: None)
```
>>> 15 <= average_error <= 25
True
```
failure_message: ❌ `average_error` is not between 15 and 25
success_message: ✅ `average_error` is correctly between 15 and 25


Tests for q5_1 (points: None)
```
>>> type(fruits) == Table
True
```
failure_message: ❌ You did not create a Table object
success_message: ✅ You created a Table and assigned it to `fruits`
```
>>> fruits.labels == ('fruit name', 'count')
True
```
failure_message: ❌ Your Table has incorrect column labels or they are in the incorrect order
success_message: ✅ Your Table has the correct column labels in the correct order


Tests for q5_2 (points: None)
```
>>> inventory.sort(0)
box ID | fruit name | count
25274  | apple      | 20
26187  | strawberry | 255
43566  | peach      | 40
48800  | orange     | 35
52357  | strawberry | 102
53686  | kiwi       | 45
57181  | strawberry | 123
57930  | grape      | 517
```
failure_message: ❌ You did not correctly load `inventory.csv`
success_message: ✅ You loaded `inventory.csv` correctly


Tests for q5_3 (points: None)
```
>>> sales.sort(0)
box ID | fruit name | count sold | price per fruit ($)
25274  | apple      | 0          | 0.8
26187  | strawberry | 25         | 0.15
43566  | peach      | 17         | 0.8
48800  | orange     | 35         | 0.6
52357  | strawberry | 102        | 0.25
53686  | kiwi       | 3          | 0.5
57181  | strawberry | 101        | 0.2
57930  | grape      | 355        | 0.06
```
failure_message: ❌ You did not correctly load `sales.csv`
success_message: ✅ You loaded `sales.csv` correctly


Tests for q5_4 (points: None)
```
>>> type(total_fruits_sold) == np.int64 or type(total_fruits_sold) == int
True
```
failure_message: ❌ Your answer should be an integer
success_message: ✅ You provided an integer answer
```
>>> total_fruits_sold > 500
True
```
failure_message: ❌ Your number is not large enough to be the total number of pieces of fruits sold
success_message: ✅ Your number is large enough to be reasonable as an answer


Tests for q5_5 (points: None)
```
>>> 50 <= total_revenue <= 150
True
```
failure_message: ❌ If you're stuck, here's a hint. You want to multiply the count sold in each box by the per-item price of fruits in that box, and then sum up each product.
success_message: ✅ Your number is between 50 and 150 dollars, which means it is a reasonable answer


Tests for q5_6 (points: None)
```
>>> remaining_inventory.num_columns
3
```
failure_message: ❌ Your table doesn't have all 3 columns
success_message: ✅ Your table has the correct number of columns
```
>>> remaining_inventory.column('count').item(0) != 45
True
```
failure_message: ❌ You forgot to subtract off the sales
success_message: ✅ You remembered to subtract off the sales
```
>>> np.array_equal(remaining_inventory.sort(0).column('fruit name'), ['apple', 'strawberry', 'peach', 'orange', 'strawberry', 'kiwi', 'strawberry', 'grape'])
True
```
failure_message: ❌ You forgot to subtract off the sales
