```python
# Initialize Otter
import otter
grader = otter.Notebook("lab06.ipynb")
```

## Lab 06: Examining the Therapeutic Touch

Welcome to Lab 6.

After such an extensive introduction to programming for data science, we are finally moving into the section of the course where we can apply our new skils to answer real questions.  

In this lab, we'll use testing techniques that were introduced in lecture to test the idea of the therapeutic touch, the idea that some practitioner can feel and massage your human energy field.

For all problems that you must write our explanations and sentences for, you **must** provide your answer in the designated space. Also, throughout this lab please be sure to not re-assign variables throughout the notebook. For example, if you use `expected_proportion_correct` in your answer to one question, do not reassign it later on. Moreover, be mindful that functions that you write in earlier cells will need to be referenced to answer questions in later questions.

```python
# Run this cell, but please don't change it.

# These lines import the Numpy and Datascience modules.
import numpy as np
from datascience import *

# These lines do some fancy plotting capabilities
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plots
plots.style.use('fivethirtyeight')
from matplotlib import patches
from ipywidgets import interact, interactive, fixed
import ipywidgets as widgets
```

## 1. What is the Therapeutic Touch

The Therapeutic Touch (TT) is the idea that everyone can feel the Human Energy Field (HEF) around individuals.  Those who practice TT have described different people's HEFs as "warm as Jell-O" and "tactile as taffy." 

TT was a popular technique used throughout the 20th century that was toted as a great way to bring balance to a person's health. Certain practitioners claim they have the ability to feel the HEF and can massage it in order to promote health and relaxation in individuals. The video below gives a brief overview.

```python
from IPython.display import YouTubeVideo
YouTubeVideo('KkXbUvU9HuI')
```

### Emily Rosa

[Emily Rosa](https://en.wikipedia.org/wiki/Emily_Rosa) was a 4th grade student who was very familiar with the world of TT, thanks to her parents, who were both medical practitioners and skeptics of TT.

For her 4th grade science fair project, Emily decided to test whether or not TT practitioners could truly interact with a person's HEF. She later went on to publish her work in TT, becoming the youngest person to have a research paper published in a peer reviewed medical journal.

### Emily's Experiment

Emily's experiment was clean, simple, and effective. Due to her parents' occupations in the medical field, she had wide access to people who claimed to be TT practitioners. 

Emily took 21 TT practitioners and used them for her science experiment. She would take a TT practitioner and ask them to extend their hands through a screen (which they can't see through). Emily would randomly extend either her left hand or her right hand (using a fair coin flip to ensure it was actually random). The TT practitioner would then have to answer which hand Emily put out. If a practitioner could truly interact with a person's HEF, it would be expected that they answered correctly.

Overall, through 210 samples, Emily observed that the practitioners picked the correct hand **44%** of the time. 

Emily's main goal here was to test whether or not the TT practitioners' guesses performed **differently** than random guesses, not better or worse. In most medical experiments, this is the norm. We want to test whether or not the treatment has an effect, **not** whether or not the treatment actually works. 

We will now begin to formulate this experiment in terms of the terminology we learned in this course. 

<!-- BEGIN QUESTION -->

### Question 1.1.

Describe the assumptions Emily made to create her model for how likely the TT practitioners are to choose the correct hand. What would be her alternative model? Discuss with students in your group to come to a conclusion. Check in with the instructor if you are stuck.


_Type your answer here, replacing this text._

<!-- END QUESTION -->

### Question 1.2.

Remember that the practitioner got the correct answer 44% (0.44) of the time. 

**According to Emily's model**, on average, what proportion of times do we expect the practitioner to guess the correct hand? Make sure your answer is a decimal between 0 and 1. 


```python
expected_proportion_correct = ...
expected_proportion_correct
```

```python
grader.check("q1_2")
```

The goal now is to see if our deviation from this expected proportion of correct answers is due to something other than chance. 

### Question 1.3.

We usually use a statistic to help determine which model the evidence points towards. What is a statistic that we can use to compare outcomes under Emily’s model to what was observed? Assign `valid_stat` to an array of integer(s) representing test statistics that Emily can use: 

1. The difference between the expected percent correct and the actual percent correct
2. The absolute difference between the expected percent correct and the actual percent correct
3. The sum of the expected percent correct and the actual percent correct

**Hint:** Format your answer like: `valid_stat = make_array(choose one of the options listed above)`


```python
valid_stat = ...
valid_stat
```

```python
grader.check("q1_3")
```

<!-- BEGIN QUESTION -->

### Question 1.4.

Why is the statistic from Question 1.3 the best choice for comparing outcomes in Emily's experiment? How does it relate to the models you defined in Question 1.1?


_Type your answer here, replacing this text._

<!-- END QUESTION -->

### Question 1.5.

Define the function `statistic` which takes in two arguments: an expected proportion of correct guesses and an observed proportion of correct guesses collected during an experiment, and returns the value of the statistic chosen in Question 1.3. Assume that the argument takes in **proportions**, but returns your answer as a **percentage**. Make sure that your function returns a value between 0 and 100!


```python
def statistic(expected_prop, observed_prop):
    ...
```

```python
grader.check("q1_5")
```

### Question 1.6.

Use your `statistic` function from Question 1.5 to calculate the observed statistic from Emily's experiment. You should know what the expected proportion should be and the observed proportion the practitioners guessed correctly based on information in previous questions.


```python
observed_statistic = ...
observed_statistic
```

```python
grader.check("q1_6")
```

### Is this observed statistic consistent with what we might see under Emily’s model?

In order to answer this question, we must simulate the experiment many times using the assumptions laid out in Emily's model. If we calculate our statistic for every simulation, we can get a feel for what values of the statistic would be expected, and which almost never occur. We can then compare our observed statistic to our collection of simulated statistics to see how consistent they are.

To help out, we'll learn about a new function to speed things up.

### **`sample_proportions`**

`sample_proportions` can be used to randomly sample from multiple categories when you know the proportion of data points that are **expected** to fall in each category. `sample_proportions` takes two arguments: the sample size and an array that contains the distribution of categories in the population (this array should sum to 1).

Consider flipping a fair coin, where the two outcomes (coin lands heads and coin lands tails) occur with an equal chance. We expect that half of all coin flips will land heads, and half of all coin flips will land tails.

Run the following cell to see the simulation of 10 flips of a fair coin. Let the first item of `coin_proportions` be the proportion of heads and the second item of `coin_proportions` be the proportion of tails.

```python
expected_proportion_heads = 0.5
expected_proportion_tails = 0.5
coin_proportions = make_array(expected_proportion_heads, expected_proportion_tails) 
ten_flips = sample_proportions(10, coin_proportions)
ten_flips
```

`sample_proportions` returns an array that is the same length as the proportion array that is passed through. It contains the proportion of each category that appears in the sample. 

In our example, the first item of `ten_flips` is the simulated proportion of heads and the second item of `ten_flips` is the simulated proportion of tails.

```python
simluated_proportion_heads = ten_flips.item(0)
simluated_proportion_tails = ten_flips.item(1)
print("In our simluation, " + str(simluated_proportion_heads) + " of flips were heads and " \
      + str(simluated_proportion_tails) + " of flips were tails.")
```

### Question 1.7.

To begin simulating, we should start by creating a representation of Emily's model to use for our simulation. This will be an array with two items in it. The first item should be the proportion of times, assuming that Emily’s model was correct, a TT practictioner picks the correct hand. The second item should be the proportion of times, under the same assumption, that the TT practitioner picks the incorrect hand. Assign `model_proportions` to this array. 

After this, we can simulate 210 hand choices, as Emily evaluated in real life, and find a single statistic to summarize this instance of the simulation. Use the `sample_proportions` function and assign the proportion of correct hand choices (out of 210) to `simulation_proportion_correct`. Lastly, use the `statistic` function that you wrote for Question 1.5 to assign `one_statistic`  to the value of the statistic for this one simulation.

**Hint:** `sample_proportions` usage can be found [here](http://data8.org/su19/python-reference.html).


```python
expected_proportion_correct = ...
expected_proportion_incorrect = ...
model_proportions = ...
simulation_proportion_correct = ...
one_statistic = ...
one_statistic
```

```python
grader.check("q1_7")
```

### Question 1.8.

Let's now see what the distribution of statistics is actually like under Emily's model. 

Define the function `simulation_and_statistic` which will compute a single simulated statistic. To do so, the `simulation_and_statistic` function will require two inputs: the expected **proportion** (a float between 0 and 1) that a TT practitioner would guess a hand correctly under a particular model which will be input as `expected_proportion_correct`, and an array that contains the proportions assumed will be correct and incorrect for the current model, input as `model_proportions`. The function should run one simulation which consists of Emily running 210 tests. The function should return the statistic of this one simulation. 

```python
def simulation_and_statistic(expected_proportion_correct, model_proportions):  
    '''Simulates 210 TT hand choices under Emily’s model.
    Returns one statistic from the simulation.'''
    ...
```

```python
grader.check("q1_8")
```

### Question 1.9.

Using the `simulation_and_statistics` function from Question 1.8, assign `simulated_statistics` to an array of 1000 statistics that you calculated under the assumption that Emily's model was true.


```python
num_repetitions = 1000
simulated_statistics = ...
for ... in ...:
    ...
```

```python
grader.check("q1_9")
```

Let's view the distribution of the simulated statistics under Emily's model, and visually compare where the observed statistic lies relative to the simulated statistics.

```python
t = Table().with_column('Simulated Statistics', simulated_statistics)
t.hist()
plots.scatter(observed_statistic, -0.01, color = 'red', s = 100);
```

We can make a visual argument as to whether we believe the observed statistic is consistent with Emily’s model. Here, since larger values of the test statistic suggest the alternative model (where the chance of guessing the correct hand is something other than 50%), we can formalize our analysis by finding what proportion of simulated statistics were as large or larger than our observed test statistic (the area at or to the right of the observed test statistic). If this area is small enough, we’ll declare that the observed data are inconsistent with our simulated model.

### Question 1.10.

Calculate the proportion of simulated statistics greater than or equal to the observed statistic. 


```python
proportion_greater_or_equal = ...
proportion_greater_or_equal
```

```python
grader.check("q1_10")
```

By convention, we often compare the proportion we just calculated to 0.05. If the proportion of simulated statistics greater than or equal to the observed statistic is sufficiently small (less than or equal to 0.05), then this is evidence against Emily's model. Otherwise, we don’t have any reason to doubt Emily’s model. 

This should help you make your own conclusions about Emily Rosa's experiment. 

<!-- BEGIN QUESTION -->

### Question 1.11.

Does the proportion calculated in Question 1.10 suggest that there is sufficient evidence to reject Emily's model (that TT practitioners are effectively guessing)?


_Type your answer here, replacing this text._

<!-- END QUESTION -->

Therapeutic touch fell out of use after this experiment, which was eventually accepted into one of the premier medical journals. TT practitioners responded and accused Emily and her family of tampering with the results, while some claimed that Emily's bad spiritual mood towards therapeutic touch made it difficult to read her HEF. Whatever it may be, Emily's experiment is a classic example about how anyone, with the right resources, can test anything they want.

# Submitting your work
You're done with this assignment! Assignments should be turned in using the following best practices:
1. Save your notebook.
2. Restart the kernel and run all cells up to this one.
3. Run the cell below with the code `grader.export(...)`. This will re-run all the tests. Make sure they are passing as you expect them to.
4. Download the file named `lab06_<date-time-stamp>.zip`, found in the explorer pane on the left side of the screen. **Note**: Clicking on the link in this notebook may result in an error, it's best to download from the file explorer panel.
5. Upload `lab06_<date-time-stamp>.zip` to the corresponding assignment on Canvas.

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit.

```python
grader.export(pdf=False, force_save=True)
```

Tests for q1_10 (points: None)
```
>>> 0 <= proportion_greater_or_equal <= 1
True
```
```
>>> proportion_greater_or_equal * 1000 == np.count_nonzero(simulated_statistics >= observed_statistic)
True
```


Tests for q1_2 (points: None)
```
>>> 0 <= expected_proportion_correct <= 1
True
```
```
>>> expected_proportion_correct == 0.5
True
```


Tests for q1_3 (points: None)
```
>>> type(valid_stat) == np.ndarray
True
```
```
>>> int(sum(valid_stat))
2
```


Tests for q1_5 (points: None)
```
>>> int(round(statistic(0.5, 0.5) + statistic(0.4, 0.1), 1))
30
```
```
>>> int(statistic(0.4, 0.1) - statistic(0.1, 0.4))
0
```


Tests for q1_6 (points: None)
```
>>> int(round(observed_statistic, 2))
6
```


Tests for q1_7 (points: None)
```
>>> expected_proportion_correct == expected_proportion_incorrect
True
```
```
>>> len(np.unique(model_proportions)) == 1
True
```
```
>>> sum(model_proportions) == 1
True
```
```
>>> type(simulation_proportion_correct) in [float, np.float64]
True
```
```
>>> 0 <= one_statistic <= 25
True
```


Tests for q1_8 (points: None)
```
>>> type(simulation_and_statistic(0.5, model_proportions)) in [float, np.float64]
True
```
```
>>> 0 <= simulation_and_statistic(0.5, model_proportions) <= 25
True
```


Tests for q1_9 (points: None)
```
>>> len(simulated_statistics) == 1000
True
```
```
>>> np.all(simulated_statistics <= 30)
True
```
```
>>> np.all(simulated_statistics >= 0)
True
```
```
>>> 0 <= simulation_and_statistic(expected_proportion_correct, model_proportions) <= 25
True
```
