```python
# Initialize Otter
import otter
grader = otter.Notebook("lab01.ipynb")
```

# Lab 01: Expressions

Welcome to MA4110: *Foundations of Data Science*!  Each week you will complete a lab assignment like this one.  You can't learn technical subjects without hands-on practice, so labs are an important part of the course.

Collaborating on labs is more than okay -- it's encouraged! You should rarely remain stuck for more than a few minutes on questions in labs, so ask a neighbor or an instructor for help. Explaining things is beneficial, too -- the best way to solidify your knowledge of a subject is to explain it. You should **not** _just_ copy/paste someone else's code, but rather work together to gain understanding of the task you need to complete. 

## Today's lab

In today's lab, you'll learn how to:

1. navigate Jupyter notebooks (like this one);
2. write and evaluate some basic *expressions* in Python, the computer language of the course;
3. call *functions* to use code other people have written; and
4. break down Python code into smaller parts to understand it.

This lab covers parts of [Chapter 3](http://www.inferentialthinking.com/chapters/03/programming-in-python.html) of the online textbook. You should read the examples in the book, but not right now. Instead, let's get started!

# 1. Jupyter notebooks
This webpage is called a Jupyter notebook. A notebook is a place to write programs and view their results, and also to write text.

## 1.1. Text cells
In a notebook, each rectangle containing text or code is called a *cell*.

Text cells (like this one) can be edited by double-clicking on them. They're written in a simple format called [Markdown](https://www.markdownguide.org/cheat-sheet/) to add formatting and section headings.  You don't need to learn Markdown, but you might want to.

After you edit a text cell, click the "run cell" button at the top menu bar that looks like ▶ or pressing `shift` + `return` on your keyboard to run the cell and advance to the next cell. (Try not to delete the instructions of the lab.)

<!-- BEGIN QUESTION -->

#### Question 1.1.1.

This paragraph is in its own text cell. Try editing it so that this sentence is the last sentence in the paragraph, and then click the "run cell" ▶ button above in the menu bar or press `shift` + `return` on your keyboard. This sentence, for example, should be deleted. So should this one.

<!-- END QUESTION -->

## 1.2. Code cells
Other cells contain code in the Python 3 language. Running a code cell will execute all of the code it contains.

To run the code in a code cell, first click on that cell to activate it.  It'll be highlighted with a little green or blue rectangle.  Next, either press ▶| or hold down `shift` + `return`.

Try running this cell:

```python
print("Hello, World!")
```

And this one:

```python
print("\N{WAVING HAND SIGN}, \N{EARTH GLOBE ASIA-AUSTRALIA}!")
```

<!-- BEGIN QUESTION -->

#### Question 1.2.1.
The fundamental building block of Python code is an expression. Cells can contain multiple lines with multiple expressions. When you run a cell, the lines of code are executed in the order in which they appear. Every `print` expression prints a line.

Change the code cell below so that it prints out:

    First this line,
    then the whole 🌏,
    and then this one.

*Hint:* If you're stuck on the Earth symbol for more than a few minutes, try talking to a neighbor or a staff member. That's a good idea for any lab problem.

```python
print('This is line one')
print('and this is line two.')
```

<!-- END QUESTION -->

## 1.3. Writing Jupyter notebooks
You can use Jupyter notebooks for your own projects or documents.  When you make your own notebook, you'll need to create your own cells for text and code.

To add a cell, click the + button in the menu bar.  It'll start out as a code cell.  You can change it to a text cell by clicking inside it so it's highlighted, clicking the drop-down box next to the restart (⟳) button in the menu bar, and choosing "Markdown".

#### Question 1.3.1.
Add a code cell below this one.  Write code in it that prints out:
   
    A whole new cell! ♪🌏♪

(That musical note symbol is like the Earth symbol.  Its long-form name is `\N{EIGHTH NOTE}`.)

Run your cell to verify that it works.

## 1.4. Errors
Python is a language, and like natural human languages, it has rules.  It differs from natural language in two important ways:
1. The rules are *simple*.  You can learn most of them in a few weeks and gain reasonable proficiency with the language in a semester.
2. The rules are *rigid*.  If you're proficient in a natural language, you can understand a non-proficient speaker, glossing over small mistakes.  A computer running Python code is not smart enough to do that.

Whenever you write code, you'll make mistakes.  When you run a code cell that has errors, Python will sometimes produce error messages to tell you what you did wrong.

Errors are okay; even experienced programmers make many errors.  When you make an error, you just have to find the source of the problem, fix it, and move on.

In the cell below, delete the last character, `)`, and then run the cell to see an error message.

```python
print("This line is missing something.")
```

You should see something like this (minus our annotations):

<img src="error.jpg" />

The last line of the error output attempts to tell you what went wrong.  The *syntax* of a language is its structure, and this `SyntaxError` tells you that you have created an illegal structure.  "`EOF`" means "end of file," so the message is saying Python expected you to write something more (in this case, a right parenthesis) before finishing the cell.

There's a lot of terminology in programming languages, but you don't need to know it all in order to program effectively. If you see a cryptic message like this, you can often get by without deciphering it.  (Of course, if you're frustrated, ask a neighbor or a staff member for help.)


#### Question 1.4.1
Fix the code above so that you can run the cell without receiving an error message.

## 1.5. The Kernel
The kernel is a program that executes the code inside your notebook and outputs the results. In the top right of your window, you can see a circle that indicates the status of your kernel. If the circle is empty (⚪), the kernel is idle and ready to execute code. If the circle is filled in (⚫), the kernel is busy running some code. 

Next to every code cell, you'll see some square brackets that look like: `[ ]`. Before you run the cell, you'll see `[ ]`. When the cell is running, you'll see `[*]`. If you see an asterisk (\*) next to a cell that doesn't go away, it's likely that the code inside the cell is taking too long to run, and it might be a good time to interrupt the kernel (discussed below). When a cell is finished running, you'll see a number inside the brackets, like so: `[1]`. The number corresponds to the order in which you run the cells; so, the first cell you run will show a 1 when it's finished running, the second will show a 2, and so on. 

You may run into problems where your kernel is stuck for an excessive amount of time, your notebook is very slow and unresponsive, or your kernel loses its connection. If this happens, try the following steps:
1. At the top of your screen, click **Kernel**, then **Interrupt**. You can also click the ⬛ button in the menu bar. Then, double check your code and try running the cell again.
2. If that doesn't help, click **Kernel**, then **Restart**. If you do this, you will have to run your code cells from the start of your notebook up until where you paused your work.
3. If that doesn't help, restart your server. This is very rarely needed, but it's good to know how to do just in case! First, save your work by clicking **File** at the top left of your screen, then **Save**. Next, click **File** again and choose **Hub Control Panel** and select **Stop my Server** to shut it down. Then, after a few moments select **Start My Server** to start it back up. Then, navigate back to the notebook you were working on. You'll still have to run your code cells again.

# 2. Numbers

Quantitative information arises everywhere in data science. In addition to representing commands to print out lines, expressions can represent numbers and methods of combining numbers. The expression `3.2500` evaluates to the number 3.25. (Run the cell and see.)

```python
3.2500
```

Notice that we didn't have to use the `print` function in the code cell above. When you run a notebook cell, if the last line has a value, then Jupyter helpfully prints out that value for you automatically. However, it won't print out values generated in the other, earlier lines of code unless you specifically ask the computer to do so using the `print` command.

```python
print(2)
3
4
```

Above, you should see that 4 is the value of the last expression and is printed to the screen.  2 is printed because the code used the `print` command. 3 is non shown when the code cell is run because it was neither printed nor was it on the last line of code in the cell.

When writing large blocks of code, you don't normally want every command printed to the screen, especially if there are thousands of values generated. Use the `print` command to only print values that you think would be helpful to see while the code is running.

## 2.1. Arithmetic
The line in the next cell subtracts.  Its value is what you'd expect.  Run it.

```python
3.25 - 1.5
```

Many basic arithmetic operations are built into Python.  The textbook section on [Expressions](http://www.inferentialthinking.com/chapters/03/1/expressions.html) describes all the arithmetic operators used in the course.  The common operator that differs from typical math notation is `**`, which raises one number to the power of the other. So, `2**3` stands for $2^3$ and evaluates to 8. 

The order of operations is the same as what you learned in elementary school, and Python also has parentheses.  For example, compare the outputs of the cells below. The second cell uses parentheses for a happy new year!

```python
7+6*5-6*3**2*2**3/4*7
```

```python
7+(6*5-(6*3))**2*((2**3)/4*7)
```

In standard math notation, the first expression is

$$7 + 6 \times 5 - 6 \times 3^2 \times \frac{2^3}{4} \times 7,$$

while the second expression is

$$7 + (6 \times 5 - \left(6 \times 3)\right)^2 \times \left(\frac{(2^3)}{4} \times 7\right).$$

<!-- BEGIN QUESTION -->

#### Question 2.1.1.
Write a Python expression in this next cell that's equal to $5 \times (3 \frac{10}{11}) - 50 \frac{1}{3} + 2^{.5 \times 22} - \frac{7}{33} + 6$.  That's five times three and ten elevenths, minus fifty and a third, plus two to the power of half twenty-two, minus seven thirty-thirds plus three.  By "$3 \frac{10}{11}$" we mean $3+\frac{10}{11}$, not $3 \times \frac{10}{11}$.

Replace the ellipses (`...`) with your expression.  Try to use parentheses only when necessary.

*Hint:* The correct output should be a familiar number.

```python
...
```

<!-- END QUESTION -->

# 3. Names
In natural language, we have terminology that lets us quickly reference very complicated concepts.  We don't say, "That's a large mammal with brown fur and sharp teeth!"  Instead, we just say, "Bear!"

In Python, we do this with *assignment statements*. An assignment statement has a name on the left side of an `=` sign and an expression to be evaluated on the right.

```python
ten = 3 * 2 + 4
```

When you run that cell, Python first computes the value of the expression on the right-hand side, `3 * 2 + 4`, which is the number 10.  Then it assigns that value to the name `ten`.  At that point, the code in the cell is done running.

After you run that cell, the value 10 is bound to the name `ten`:

```python
ten
```

The statement `ten = 3 * 2 + 4` is not asserting that `ten` is already equal to `3 * 2 + 4`, as we might expect by analogy with math notation.  Rather, that line of code changes what `ten` means; it now refers to the value 10, whereas before it meant nothing at all.

If the designers of Python had been ruthlessly pedantic, they might have made us write

    define the name ten to hereafter have the value of 3 * 2 + 4 

instead.  You will probably appreciate the brevity of "`=`"!  But keep in mind that this is the real meaning.

<!-- BEGIN QUESTION -->

#### Question 3.0.1.
Run the following cell which uses a variable name `eleven` that hasn't been assigned to anything. You'll see an error! Correct the code in the cell below so when it is run it does not generate an error message.

```python
eleven + 8
```

<!-- END QUESTION -->

A common pattern in Jupyter notebooks is to assign a value to a name and then immediately evaluate the name in the last line in the cell so that the value is displayed as output. 

```python
close_to_pi = 355/113
close_to_pi
```

Another common pattern is that a series of lines in a single cell will build up a complex computation in stages, naming the intermediate results.

```python
semimonthly_salary = 842.5
monthly_salary = 2 * semimonthly_salary
number_of_months_in_a_year = 12
yearly_salary = number_of_months_in_a_year * monthly_salary
yearly_salary
```

Names in Python can have letters (upper- and lower-case letters are both okay and count as different letters), underscores, and numbers.  The first character can't be a number (otherwise a name might look like a number).  And names can't contain spaces, since spaces are used to separate pieces of code from each other.

Other than those rules, what you name something doesn't matter *to Python*.  For example, this cell does the same thing as the above cell, except everything has a different name:

```python
a = 842.5
b = 2 * a
c = 12
d = c * b
d
```

**However**, names are very important for making your code *readable* to yourself and others.  The cell above is shorter, but it's totally useless without an explanation of what it does.

## 3.1. Checking your code
Now that you know how to name things, you can start using the built-in *tests* to check whether your work is correct. Sometimes, there are multiple tests for a single question, and passing all of them is required to receive credit for the question. A test cell will always look like:

```
grader.check("q1_2_1")
```

Go ahead and attempt Question 3.1.1. Running the test cell directly after it will test whether you have assigned `seconds_in_a_decade` correctly in Question 3.1.1. If you haven't, this particular test will tell you the correct answer. Resist the urge to just copy it, and instead try to adjust your expression. Additionally, if you make a common error sometimes the tests will give hints about what went wrong.

#### Question 3.1.1.
Assign the name `seconds_in_a_decade` to the number of seconds between midnight January 1, 2010 and midnight January 1, 2020. Note that there are two leap years in this span of a decade. A non-leap year has 365 days and a leap year has 366 days.

If you run the test cell below, and receive the error:

`NameError: name 'grader' is not defined`

That means you didn't run the very first code cell in this notebook! Scroll back up to the top of your notebook, run the very first cell, and then try running the test cell again.

```python
seconds_in_a_decade = ...

# We've put this line in this cell  so that it will print the value you've given 
# to seconds_in_a_decade when you run it. You don't need to change this.
seconds_in_a_decade
```

```python
grader.check("q3_1_1")
```

# 3.2. Comments
You may have noticed these lines in the cell in which you answered Question 3.2:

    # We've put this line in this cell  so that it will print the value you've given 
    # to seconds_in_a_decade when you run it. You don't need to change this.
    
This is called a *comment*. It doesn't make anything happen in Python; Python ignores anything on a line after a `#`.  Instead, it's there to communicate something about the code to you, the human reader. Comments are extremely useful. 


<img src="http://imgs.xkcd.com/comics/future_self.png">

Source: [XKCD](http://imgs.xkcd.com/comics/future_self.png)

## 3.3. Application: A physics experiment

On the Apollo 15 mission to the Moon, astronaut David Scott famously replicated Galileo's physics experiment in which he showed that gravity accelerates objects of different mass at the same rate. Because there is no air resistance for a falling object on the surface of the Moon, even two objects with very different masses and densities should fall at the same rate. David Scott compared a feather and a hammer.

You can run the following cell to watch a video of the experiment.

```python
from IPython.display import YouTubeVideo
# The original URL is: 
# https://www.youtube.com/watch?v=U7db6ZeLR5s

YouTubeVideo("U7db6ZeLR5s")
```

Here's the transcript of the video:

**167:22:06 Scott**: Well, in my left hand, I have a feather; in my right hand, a hammer. And I guess one of the reasons we got here today was because of a gentleman named Galileo, a long time ago, who made a rather significant discovery about falling objects in gravity fields. And we thought where would be a better place to confirm his findings than on the Moon. And so we thought we'd try it here for you. The feather happens to be, appropriately, a falcon feather for our Falcon. And I'll drop the two of them here and, hopefully, they'll hit the ground at the same time. 

**167:22:43 Scott**: How about that!

**167:22:45 Allen**: How about that! (Applause in Houston)

**167:22:46 Scott**: Which proves that Mr. Galileo was correct in his findings.

**Newton's Law.** Using this footage, we can also attempt to confirm another famous bit of physics: Newton's law of universal gravitation. Newton's laws predict that any object dropped near the surface of the Moon should fall

$$\frac{1}{2} G \frac{M}{R^2} t^2 \text{ meters}$$

after $t$ seconds, where $G$ is a universal constant, $M$ is the moon's mass in kilograms, and $R$ is the moon's radius in meters.  So if we know $G$, $M$, and $R$, then Newton's laws let us predict how far an object will fall over any amount of time.

To verify the accuracy of this law, we will calculate the difference between the predicted distance the hammer drops and the actual distance.  (If they are different, it might be because Newton's laws are wrong, or because our measurements are imprecise, or because there are other factors affecting the hammer for which we haven't accounted.)

#### Question 3.3.1.
Someone studied the video and estimated that the hammer was dropped 113 cm from the surface. Counting frames in the video, the hammer falls for 1.2 seconds (36 frames). Complete the code in the next cell to fill in the data from the experiment.

*Hint:* No computation required; just fill in data from the paragraph above.

```python
# Fill in the time, measured in seconds
time = ...

# Fill in the estimated distance the hammer fell, in meters.
estimated_distance_m = ...
```

```python
grader.check("q3_3_1")
```

#### Question 3.3.2.
Now, complete the code in the next cell to compute the difference between the predicted and estimated distances (in meters) that the hammer fell in this experiment.

This just means translating the formula above ($\frac{1}{2}G\frac{M}{R^2}t^2$) into Python code.  You'll have to replace each variable in the math formula with the name we gave that number in Python code.

*Hint:* Try to use variables you've already defined in the previous question!

```python
# First, we've written down the values of the 3 universal constants 
# that show up in Newton's formula.

# G, the universal constant measuring the strength of gravity.
gravity_constant = 6.674 * 10**-11

# M, the moon's mass, in kilograms.
moon_mass_kg = 7.34767309 * 10**22

# R, the radius of the moon, in meters.
moon_radius_m = 1.737 * 10**6
```

Compute the distance the hammer should have fallen over the duration of the fall, in meters, according to Newton's law of gravity. The text above describes the formula for this distance given by Newton's law. Use the values assigned in the cells above as part of your calculation.

```python
predicted_distance_m = ...
predicted_distance_m
```

Here we've computed the difference between the predicted fall distance and the distance we actually measured. If you've filled in the above code, this should just work.

```python
difference = predicted_distance_m - estimated_distance_m
difference
```

```python
grader.check("q3_3_2")
```

# 4. Calling functions

The most common way to combine or manipulate values in Python is by calling functions. Python comes with many built-in functions that perform common operations.

For example, the `abs` function takes a single number as its argument and returns the absolute value of that number. Run the next two cells and see if you understand the output.

```python
abs(5)
```

```python
abs(-5)
```

## 4.1. Application: Computing walking distances
Chunhua is on the corner of 7th Avenue and 42nd Street in Midtown Manhattan, and she wants to know far she'd have to walk to get to Gramercy School on the corner of 10th Avenue and 34th Street.

She can't cut across blocks diagonally, since there are buildings in the way.  She has to walk along the sidewalks.  Using the map below, she sees she'd have to walk 3 avenues (long blocks) and 8 streets (short blocks).  In terms of the given numbers, she computed 3 as the difference between 7 and 10, *in absolute value*, and 8 similarly.  

Chunhua also knows that blocks in Manhattan are all about 80m by 274m (avenues are farther apart than streets).  So in total, she'd have to walk $(80 \times |42 - 34| + 274 \times |7 - 10|)$ meters to get to the park.

<img src="map.jpg"/>

#### Question 4.1.1.
Fill in the line `num_avenues_away = ...` in the next cell so that the cell calculates the distance Chunhua must walk and gives it the name `manhattan_distance`.  Everything else has been filled in for you.  **Use the `abs` function.** Also, be sure to run the test cell afterward to test your code.

```python
# Here's the number of streets away:
num_streets_away = abs(42-34)

# Compute the number of avenues away in a similar way:
num_avenues_away = ...

street_length_m = 80
avenue_length_m = 274

# Now we compute the total distance Chunhua must walk.
manhattan_distance = street_length_m*num_streets_away + avenue_length_m*num_avenues_away

# We've included this line so that you see the distance you've computed 
# when you run this cell.  
# You don't need to change it, but you can if you want.
manhattan_distance
```

```python
grader.check("q4_1_1")
```

### Multiple arguments

Some functions take multiple arguments, separated by commas. For example, the built-in `max` function returns the maximum argument passed to it.

```python
max(2, -3, 4, -5)
```

# 5. Understanding nested expressions
Function calls and arithmetic expressions can themselves contain expressions.  You saw an example in the last question:

    abs(42-34)

has 2 number expressions in a subtraction expression in a function call expression.  And you probably wrote something like `abs(7-10)` to compute `num_avenues_away`.

Nested expressions can turn into complicated-looking code. However, the way in which complicated expressions break down is very regular.

Suppose we are interested in heights that are very unusual.  We'll say that a height is unusual to the extent that it's far away on the number line from the average human height.  [An estimate](http://press.endocrine.org/doi/full/10.1210/jcem.86.9.7875?ck=nck&) of the average adult human height (averaging, we hope, over all humans on Earth today) is 1.688 meters.

So if Kayla is 1.21 meters tall, then her height is $|1.21 - 1.688|$, or $.478$, meters away from the average.  Here's a picture of that:

<img src="numberline_0.png">

And here's how we'd write that in one line of Python code:

```python
abs(1.21 - 1.688)
```

What's going on here?  `abs` takes just one argument, so the stuff inside the parentheses is all part of that *single argument*.  Specifically, the argument is the value of the expression `1.21 - 1.688`.  The value of that expression is `-.478`.  That value is the argument to `abs`.  The absolute value of that is `.478`, so `.478` is the value of the full expression `abs(1.21 - 1.688)`.

Picture simplifying the expression in several steps:

1. `abs(1.21 - 1.688)`
2. `abs(-.478)`
3. `.478`

In fact, that's basically what Python does to compute the value of the expression.

#### Question 5.0.1.
Say that Paola's height is 1.76 meters.  In the next cell, use `abs` to compute the absolute value of the difference between Paola's height and the average human height.  Give that value the name `paola_distance_from_average_m`.

<img src="numberline_1.png">

Replace the `...` in the code cell below with an expression to compute the absolute value of the difference between Paola's height (1.76m) and the average human height.

```python
paola_distance_from_average_m = ...

# We've written this here so that the distance you 
# compute will get printed when you run this cell.
paola_distance_from_average_m
```

```python
grader.check("q5_0_1")
```

## 5.1. More nesting
Now say that we want to compute the more unusual of the two heights.  We'll use the function `max`, which (again) takes two numbers as arguments and returns the larger of the two arguments.  Combining that with the `abs` function, we can compute the larger distance from average among the two heights:

```python
# Just read and run this cell.

kayla_height_m = 1.21
paola_height_m = 1.76
average_adult_height_m = 1.688

# The larger distance from the average human height, among the two heights:
larger_distance_m = max(abs(kayla_height_m - average_adult_height_m), abs(paola_height_m - average_adult_height_m))

# Print out our results in a nice readable format:
print("The larger distance from the average height among these two people is", larger_distance_m, "meters.")
```

The line where `larger_distance_m` is computed looks complicated, but we can break it down into simpler components just like we did before.

The basic recipe is to repeatedly simplify small parts of the expression:
* **Basic expressions:** Start with expressions whose values we know, like names or numbers.
    - Examples: `paola_height_m` or `5`.
* **Find the next simplest group of expressions:** Look for basic expressions that are directly connected to each other. This can be by arithmetic or as arguments to a function call. 
    - Example: `kayla_height_m - average_adult_height_m`.
* **Evaluate that group:** Evaluate the arithmetic expression or function call. Use the value computed to replace the group of expressions.  
    - Example: `kayla_height_m - average_adult_height_m` becomes `-.478`.
* **Repeat:** Continue this process, using the value of the previously-evaluated expression as a new basic expression. Stop when we've evaluated the entire expression.
    - Example: `abs(-.478)` becomes `.478`, and `max(.478, .072)` becomes `.478`.

You can run the next cell to see a slideshow of that process.

```python
from IPython.display import IFrame
IFrame('https://docs.google.com/presentation/d/e/2PACX-1vTiIUOa9tP4pHPesrI8p2TCp8WCOJtTb3usOacQFPfkEfvQMmX-JYEW3OnBoTmQEJWAHdBP6Mvp053G/embed?start=false&loop=false&delayms=3000', 800, 600)
```

Ok, your turn. 

#### Question 5.1.1.
Given the heights of three players from the basketball team Golden State Warriors, write an expression that computes the smallest difference between any of the three heights. Your expression shouldn't have any numbers in it, only function calls and the variable names. Give the value of your expression the name `min_height_difference`.

```python
# The three players' heights, in meters:
ball =  2.01 # LaMelo Ball is 6'7"
rozier = 1.85 # Terry Rozier is 6'1"
richards = 1.13 # Nick Richards is 7'0"
             
# We'd like to look at all 3 pairs of heights, 
# compute the absolute difference between each pair, 
# and then find the smallest of those 3 absolute differences.  
# This is left to you!  
# If you're stuck, try computing the value for each step of the process 

min_height_difference = ...
min_height_difference
```

```python
grader.check("q5_1_1")
```

# 6. Submitting your work
You're done with Lab 01! All assignments in the course will be distributed as notebooks like this one. All assignments should be turned in using the following best practices:
1. Save your notebook.
2. Restart the kernel and run all cells up to this one.
3. Run the cell below with the code `grader.export(...)`. This will re-run all the tests. Make sure they are passing as you expect them to.
4. Download the file named `lab01_<date-time-stamp>.zip`, found in the explorer pane on the left side of the screen. **Note**: Clicking on the link in this notebook may result in an error, it's best to download from the file explorer panel.
5. Upload `lab01_<date-time-stamp>.zip` to the Lab 01 assignment on Canvas.

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit.

When done exporting, download the .zip file by pressing `SHIFT` on your keyboard and then right-clicking on the generated link and selecting **Save Link As**. Or, find the `.zip` file in the left side of the screen and right-click and select **Download**. You'll submit this `.zip` file for the assignment in Canvas to Gradescope for grading.

```python
grader.export(pdf=False, force_save=True)
```

Tests for q3_1_1 (points: 1)
```
>>> 'seconds_in_a_decade' in vars()
True
```
failure_message: ❌ ❌ It looks like you didn't give anything the name seconds_in_a_decade. Maybe there's a typo, or maybe you just need to run the cell below Question 3.1.1 where you defined seconds_in_a_decade.
success_message: ✅ You correctly defined a variable named seconds_in_a_decade
```
>>> seconds_in_a_decade != ...
True
```
failure_message: ❌ It looks like you didn't change the cell to define seconds_in_a_decade appropriately.  It should be a number,computed using Python's arithmetic. For example, this is almost right, seconds_in_a_decade = 10*365*24*60*60
success_message: ✅ You correctly defined a variable named seconds_in_a_decade to be something other than ...
```
>>> seconds_in_a_decade != 315360000
True
```
failure_message: ❌ It looks like you didn't account for leap years. There were 2 leap years and 8 non-leap years in this period of time. Leap years have 366 days instead of 365.
success_message: ✅ You correctly did not compute the number of seconds in 10 non-leap years.
```
>>> seconds_in_a_decade == 315532800
True
```
failure_message: ❌ You did not compute the correct answer
success_message: ✅ You correctly computed the correct answer


Tests for q3_3_1 (points: 1)
```
>>> time != ...
True
```
failure_message: ❌ You didn't fill in the variable named time. Fill it in with a number like, time = 4.457
success_message: ✅ You provided a value for the variable named time.
```
>>> round(time, 5)
1.2
```
failure_message: ❌ You didn't provide the correct value for time. Re-read the paragraph above and try again.
success_message: ✅ You provided the correct value for the variable named time.
```
>>> estimated_distance_m != ...
True
```
failure_message: ❌ You didn't fill in the variable named estimated_distance_m. Fill it in with a number like, estimated_distance_m = 4.457
success_message: ✅ You provided a value for the variable named estimated_distance_m.
```
>>> estimated_distance_m != 113
True
```
failure_message: ❌ You provided the correct distance, but in centimeters not meters.
```
>>> round(estimated_distance_m, 5)
1.13
```
failure_message: ❌ You didn't provide the correct value for estimated_distance_m. Re-read the paragraph above and try again.
success_message: ✅ You provided the correct value for the variable named estimated_distance_m.


Tests for q3_3_2 (points: 1)
```
>>> predicted_distance_m != ...
True
```
failure_message: ❌ You didn't fill in the variable named predicted_distance_m. Fill it in with a calculation like, predicted_distance_m = rate * time
success_message: ✅ You provided a value for the variable named predicted_distance_m.
```
>>> round(predicted_distance_m, 5)
1.17022
```
failure_message: ❌ You did not correctly compute the predicted distance in meters.
success_message: ✅ You correctly computed the predicted distance in meters
```
>>> round(difference, 5)
0.04022
```
failure_message: ❌ You did not correctly compute the predicted difference in meters.
success_message: ✅ You correctly computed the predicted difference in meters


Tests for q4_1_1 (points: 1)
```
>>> num_avenues_away != ...
True
```
```
>>> num_avenues_away != -3
True
```
```
>>> num_avenues_away
3
```


Tests for q5_0_1 (points: 1)
```
>>> round(paola_distance_from_average_m, 5)
0.072
```


Tests for q5_1_1 (points: 1)
```
>>> round(min_height_difference, 5)
0.16
```
