```python
# Initialize Otter
import otter
grader = otter.Notebook("hw01.ipynb")
```

# Homework 1: Causality and Expressions

Please complete this notebook by filling in the cells provided. Before you begin, be sure you've run the very first cell in this notebook to load `otter`, the auto-grading library that will provide feedback while completing this assignment. This cell should be run each time you load this notebook, as everything stored into memory is reset when you close Datahub for more than a few minutes.

**Recommended Reading:**
- [What is Data Science](http://www.inferentialthinking.com/chapters/01/what-is-data-science.html)
- [Causality and Experiments](http://www.inferentialthinking.com/chapters/02/causality-and-experiments.html) 
- [Programming in Python](http://www.inferentialthinking.com/chapters/03/programming-in-python.html)

For all problems that you must write explanations and sentences for, you **must** provide your answer in the designated space. Moreover, throughout this homework and all future ones, please be sure to not re-assign variables throughout the notebook! For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Otherwise, you will fail tests that you thought you were passing previously!

**Note: This homework has hidden tests on it. Additional tests will be run once your homework is submitted for grading. While you may pass all the tests you have access to before submission, you may not earn full credit if you do not pass the hidden tests as well.**. 

Many of the tests you have access to before submitting only test to ensure you have given an answer that is formatted correctly and/or you have given an answer that *could* make sense in context. For example, a test you have access to while completing the assignment may check that you selected a valid choice for a multiple choice problem (1, 2, or 3) or that your answer is an integer between 0 and 50 if asked to count a subset of states in the United States. The tests that are run after submission will evaluate your work for accuracy. **Do not assume that just because all your tests pass before submission means that your answers are correct!**

Consult with your teacher and course syllabus for information and policies regarding appropriate collaboration with other students, appropriate use of AI tools, and submission of late work.

## 1. Scary Arithmetic

An ad for ADT Security Systems says,

> "When you go on vacation, burglars go to work [...] According to FBI statistics, over 25% of home burglaries occur between Memorial Day to Labor Day."

<!-- BEGIN QUESTION -->

#### Question 1.1.
This ad is suggesting that you should be worried about going on your summer vacation because there's an increased chance that your home will be broken into during this time period. Do the data in the ad support the claim that burglars are more likely to go to work during the time between Memorial Day to Labor Day than other times of the year? Put another way, should you believe the security company that you should be on alert during the summer for an increased risk of burglary? Please explain your answer.

**Note:** You can assume that "over 25%" means only slightly over. Had it been much over, say closer to 30%, then the marketers would have said so.


_Type your answer here, replacing this text._

<!-- END QUESTION -->

## 2. Characters in Little Women

We can use some Python code to count the number of times that the literary characters were named in each chapter of the classic book, [*Little Women*](https://www.inferentialthinking.com/chapters/01/3/1/literary-characters). In computer science, the word "character" also refers to a letter, digit, space, or punctuation mark; any single element of a text. The following code generates a scatter plot in which each dot corresponds to a chapter of *Little Women*. The horizontal position of a dot measures the number of periods in the chapter. The vertical position measures the total number of characters.

```python
# This cell contains code that hasn't yet been covered in the course,
# but you should be able to interpret the scatter plot it generates.

from datascience import *
from urllib.request import urlopen
import numpy as np
%matplotlib inline

little_women_url = 'https://www.inferentialthinking.com/data/little_women.txt'
chapters = urlopen(little_women_url).read().decode().split('CHAPTER ')[1:]
text = Table().with_column('Chapters', chapters)
Table().with_columns(
    'Periods',    np.char.count(chapters, '.'),
    'Characters', text.apply(len, 0)
    ).scatter(0)
```

#### Question 2.1.
Around how many periods are there in the chapter with the most characters? Assign either 1, 2, 3, 4, or 5 to the name `characters_q1` below.

1. 250
2. 390
3. 440
4. 32,000
5. 40,000


```python
characters_q1 = ...
```

```python
grader.check("q2_1")
```

#### Question 2.2.
Which of the following chapters has the most characters per period? Assign either 1, 2, or 3 to the name `characters_q2` below.
1. The chapter with about 60 periods
2. The chapter with about 350 periods
3. The chapter with about 440 periods


```python
characters_q2 = ...
```

```python
grader.check("q2_2")
```

Again, the test above checks that your answers are in the correct format, but not that you have answered correctly.

To discover more interesting facts from this plot, read [Section 1.3.2](https://www.inferentialthinking.com/chapters/01/3/2/another-kind-of-character) of the textbook.

## 3. Names and Assignment Statements

#### Question 3.1.
If you were to run the code in the screen shot of the cell shown below, Python produces the displayed cryptic error message.

![Error message. Input is: 4 = 2 + 2  Output: SyntaxError: cannot assign to literal here. Maybe you meant '==' instead of '='?](error.png)

Choose the best explanation of what's wrong with the code, and then assign 1, 2, 3, or 4 to `names_q1` below to indicate your answer.

1. Python is smart and already knows `4 = 2 + 2`.

2. It should be `2 + 2 = 4`.

3. `4` is not a variable, so you can't assign a value to it.

4. If I create a new code cell and attempt to run this code, I don't get an error message. This is a trick question.


```python
names_q1 = ...
```

```python
grader.check("q3_1")
```

#### Question 3.2.
If you were to run the code in the screen shot of the cell shown below, Python would produce another cryptic error message.

![Error message. Input: first line: two = 3, second line: six = two plus two  Output: SyntaxError: invalid syntax](error2.png)

Choose the best explanation of what's wrong with the code and assign 1, 2, 3, or 4 to `names_q2` below to indicate your answer.

1. The `plus` operation only applies to numbers, not the word "two".

2. The name "two" cannot be assigned to the number 3.

3. Two plus two is four, not six.

4. Python cannot interpret the name `two` followed directly by a name that has not been defined.


```python
names_q2 = ...
```

```python
grader.check("q3_2")
```

#### Question 3.3.
If you were to run the code in the screen shot of the cell shown below, Python would produce yet another cryptic error message.

![Error message](error3.png)

Choose the best explanation of what's wrong with the code and assign 1, 2, or 3 to `names_q3` below to indicate your answer.

1. Python doesn't want `y` to be assigned.

2. The `print` operation is meant for displaying values to the programmer, not for assigning values!

3. Python can’t do addition between one name and one number. It has to be 2 numbers or 2 predefined names.


```python
names_q3 = ...
```

```python
grader.check("q3_3")
```

## 4. Job Opportunities & Education in Rural India

A [study](http://www.nber.org/papers/w16021.pdf) at UCLA investigated factors that might result in greater attention to the health and education of girls in rural India. One such factor is information about job opportunities for women. The idea is that if people know that educated women can get good jobs, they might take more care of the health and education of girls in their families, as an investment in the girls’ future potential as earners. Without the knowledge of job opportunities, the author hypothesizes that families do not invest in women’s well-being.

The study focused on 160 villages outside the capital of India, all with little access to information about call centers and similar organizations that offer job opportunities to women. In 80 of the villages chosen at random, recruiters visited the village, described the opportunities, recruited women who had some English language proficiency and experience with computers, and provided ongoing support free of charge for three years. In the other 80 villages, no recruiters visited and no other intervention was made.

At the end of the study period, the researchers recorded data about the school attendance and health of the children in the villages.

#### Question 4.1.
Which statement best describes the *treatment* and *control* groups for this study? Assign either 1, 2, or 3 to the name `jobs_q1` below.

1. The treatment group was the 80 villages visited by recruiters, and the control group was the other 80 villages with no intervention.

2. The treatment group was the 160 villages selected, and the control group was the rest of the villages outside the capital of India.

3. There is no clear notion of *treatment* and *control* group in this study.


```python
jobs_q1 = ...
```

```python
grader.check("q4_1")
```

#### Question 4.2.
Was this an observational study or a randomized controlled experiment? Assign either 1, 2, or 3 to the name `jobs_q2` below.

1. This was an observational study.

2. This was a randomized controlled experiment.  

3. This was a randomized observational study.


```python
jobs_q2 = ...
```

```python
grader.check("q4_2")
```

<!-- BEGIN QUESTION -->

#### Question 4.3.
The study reported, "Girls aged 5-15 in villages that received the recruiting services were 3 to 5 percentage points more likely to be in school and experienced an increase in Body Mass Index, reflecting greater nutrition and/or medical care. However, there was no net gain in height. For boys, there was no change in any of these measures." 

Why do you think the author thinks it is important to point out the lack of change in these measures in the boys, when this study is focused on improving outcomes for girls? Please fully explain your answer.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## 5. Differences in Math Placement


NCSSM’s Math Department keeps records regarding incoming junior's initial Math course. The table below displays the numbers of first semester juniors in each of the listed courses academic years 2010-2011 and 2020-2021.

| Course                             | 2010-2011    | 2020-2021   |
|------------------------------------|--------------|-------------|
| Precalculus                        |      162     |    63       |
| Calculus                           |      72      |    175      |
| Multivariable Calculus             |      18      |    53       |



#### Question 5.1. 
Suppose you want to find the **biggest** absolute difference between the numbers of students in the two years, among the three courses.

In the cell below, compute this value and call it `biggest_change`. Use a single expression (a single line of code) to compute the answer. Let Python perform all the arithmetic (like subtracting 18 from 53) rather than simplifying the expression yourself. The built-in `abs` function takes a numerical input and returns the absolute value. The built-in `max` function can take in 3 arguments and returns the maximum of the three numbers


```python
biggest_change = ...
biggest_change
```

```python
grader.check("q5_1")
```

#### Question 5.2.
Which of the three courses had the **smallest** absolute difference? Assign `smallest_change_course` to 1, 2, or 3 where each number corresponds to the following course:

1: Precalculus  
2: Calculus  
3: Multivariable Calculus

Choose the number that corresponds to the course with the smallest absolute difference.

You should be able to answer by rough mental arithmetic, without having to calculate the exact value for each course. 


```python
smallest_change_course = ...
smallest_change_course
```

```python
grader.check("q5_2")
```

#### Question 5.3.
For each course, define the “relative change” to be the following: $\large{\frac{\text{absolute difference}}{\text{value in 2010-2011}} \cdot 100}$ 

Fill in the code below such that `precalculus_relative_change`, `calculus_relative_change` and `multi_relative_change` are assigned to the relative changes for their respective courses.


```python
precalculus_relative_change = (abs(...) / 162) * 100
calculus_relative_change = ...
multi_relative_change = ...
precalculus_relative_change, calculus_relative_change, multi_relative_change
```

```python
grader.check("q5_3")
```

#### Question 5.4.
Assign `biggest_rel_change_course` to the course with the biggest relative change, where 1, 2, and 3  correspond to to the following: 

1. Precalculus 
2. Calculus  
3. Multivariable Calculus

```python
biggest_rel_change_course = ...
biggest_rel_change_course
```

```python
grader.check("q5_4")
```

## 6. Nearsightedness Study


Myopia, or nearsightedness, results from a number of genetic and environmental factors. In 1999, [Quinn et al](https://doi.org/10.1038/20094) studied the relation between myopia and ambient lighting at night (for example, from nightlights or room lights) during childhood.

<!-- BEGIN QUESTION -->

#### Question 6.1.
The data were gathered by the following procedure, reported in the study. "Between January and June 1998, parents of children aged 2-16 years [...] that were seen as outpatients in a university pediatric ophthalmology clinic completed a questionnaire on the child's light exposure both at present and before the age of 2 years." Was this study observational, or was it a controlled experiment? Explain. 


_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

#### Question 6.2.
The study found that of the children who slept with a room light on before the age of 2, 55% were myopic. Of the children who slept with a night light on before the age of 2, 34% were myopic. Of the children who slept in the dark before the age of 2, 10% were myopic. The study concluded that, "The prevalence of myopia [...] during childhood was strongly associated with ambient light exposure during sleep at night in the first two years after birth."

Do the data support this statement? You may interpret "strongly" in any reasonable qualitative way.


_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

#### Question 6.3.
On May 13, 1999, CNN reported the results of this study under the headline, "Night light may lead to nearsightedness." Does the conclusion of the study claim that the use of a night light causes nearsightedness?


_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

#### Question 6.4.
The final paragraph of the CNN report said that "several eye specialists" had pointed out that the study should have accounted for heredity.

Myopia is passed down from parents to children. Myopic parents are more likely to have myopic children, and may also be more likely to leave lights on habitually (since the parents have poor vision). In what way does the knowledge of this possible genetic link affect how we interpret the data from the study? Explain.


_Type your answer here, replacing this text._

<!-- END QUESTION -->

#### Question 6.5.
The Reverend Henry Whitehead was skeptical of John Snow's conclusion about the Broad Street pump. After the Broad Street cholera epidemic ended, Whitehead set about trying to prove Snow wrong. (The history of the event is detailed [here](https://www.ncbi.nlm.nih.gov/pmc/articles/PMC1034367/pdf/medhist00183-0026.pdf).)

He realized that Snow had focused his analysis almost entirely on those who had died. Whitehead, therefore, investigated the drinking habits of people in the Broad Street area who had not died in the outbreak.

What is the main reason it was important to study this group? Assign either 1, 2, or 3 to the name `survivor_answer` below.

1. If Whitehead had found that many people had drunk water from the Broad Street pump and not caught cholera, that would have been evidence against Snow's hypothesis.
2. Survivors could provide additional information about what else could have caused the cholera, potentially unearthing another cause.
3. Through considering the survivors, Whitehead could have identified a cure for cholera.

```python
survivor_answer = ...
```

```python
grader.check("q6_5")
```

# Submitting your work
You're done with this assignment! Assignments should be turned in using the following best practices:
1. Save your notebook.
2. Restart the kernel and run all cells up to this one.
3. Run the cell below with the code `grader.export(...)`. This will re-run all the tests. Make sure they are passing as you expect them to.
4. Download the file named `hw01_<date-time-stamp>.zip`, found in the explorer pane on the left side of the screen. **Note**: Clicking on the link in this notebook may result in an error, it's best to download from the file explorer panel.
5. Upload `hw01_<date-time-stamp>.zip` to the corresponding assignment on Canvas.

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit.

When done exporting, download the .zip file by pressing `SHIFT` on your keyboard and then right-clicking on the generated link and selecting **Save Link As**. Or, find the `.zip` file in the left side of the screen and right-click and select **Download**. You'll submit this `.zip` file for the assignment in Canvas to Gradescope for grading.

```python
grader.export(pdf=False, force_save=True)
```

Tests for q2_1 (points: None)
```
>>> characters_q1 != ...
True
```
failure_message: ❌ It looks like you didn't replace ... with an answer choice
```
>>> characters_q1 in [1, 2, 3, 4, 5]
True
```
failure_message: ❌ It looks like you didn't select 1, 2, 3, 4, or 5. Maybe you typed in the number of periods?
success_message: ✅ You selected a valid choice (1, 2, 3, 4, or 5)


Tests for q2_2 (points: None)
```
>>> characters_q2 != ...
True
```
failure_message: ❌ It looks like you didn't replace ... with an answer choice
```
>>> characters_q2 in [1, 2, 3]
True
```
failure_message: ❌ It looks like you didn't select 1, 2, or 3. Maybe you typed in the number of periods?
success_message: ✅ You selected a valid choice (1, 2, or 3)


Tests for q3_1 (points: None)
```
>>> names_q1 in [1, 2, 3, 4]
True
```
failure_message: ❌ It looks like you didn't select 1, 2, 3, or 4.
success_message: ✅ You selected a valid choice (1, 2, 3, or 4)


Tests for q3_2 (points: None)
```
>>> names_q2 in [1, 2, 3, 4]
True
```
failure_message: ❌ It looks like you didn't select 1, 2, 3, or 4.
success_message: ✅ You selected a valid choice (1, 2, 3, or 4)


Tests for q3_3 (points: None)
```
>>> names_q3 in [1, 2, 3]
True
```
failure_message: ❌ It looks like you didn't select 1, 2, or 3.
success_message: ✅ You selected a valid choice (1, 2, or 3)


Tests for q4_1 (points: None)
```
>>> jobs_q1 in [1, 2, 3]
True
```
failure_message: ❌ It looks like you didn't select 1, 2, or 3.
success_message: ✅ You selected a valid choice (1, 2, or 3)


Tests for q4_2 (points: None)
```
>>> jobs_q2 in [1, 2, 3]
True
```
failure_message: ❌ It looks like you didn't select 1, 2, or 3.
success_message: ✅ You selected a valid choice (1, 2, or 3)


Tests for q5_1 (points: None)
```
>>> import numpy as np
>>> isinstance(biggest_change, (int, float, np.float64))
True
```
failure_message: ❌ You did not provide an an answer that is an integer or a floating point number.
success_message: ✅ You have provided a numerical answer, which this question requires.


Tests for q5_2 (points: None)
```
>>> isinstance(smallest_change_course, (int, float, np.float64))
True
```
failure_message: ❌ You did not provide an an answer that is an integer or a floating point number.
success_message: ✅ You have provided a numerical answer, which this question requires.


Tests for q5_3 (points: None)
```
>>> (isinstance(precalculus_relative_change, (int, float, np.float64)), isinstance(calculus_relative_change, (int, float, np.float64)), isinstance(multi_relative_change, (int, float, np.float64))) == (True, True, True)
True
```
failure_message: ❌ At least one variable is not an integer or a floating point number.
success_message: ✅ You have provided a numerical answer for each calculation, which this question requires.
```
>>> multi_relative_change > calculus_relative_change > precalculus_relative_change > 50
True
```
failure_message: ❌ Each of your values should be a percentage greater than 50. Double check you are computing percents, and not proportions.
success_message: ✅ You correctly computed the values as percents, and not as decimals


Tests for q5_4 (points: None)
```
>>> biggest_rel_change_course in [1, 2, 3]
True
```
failure_message: ❌ It looks like you didn't select 1, 2, or 3.
success_message: ✅ You selected a valid choice (1, 2, or 3)


Tests for q6_5 (points: None)
```
>>> survivor_answer in [1, 2, 3]
True
```
failure_message: ❌ It looks like you didn't select 1, 2, or 3.
success_message: ✅ You selected a valid choice (1, 2, or 3)
